```python
from numpy.testing import assert_equal
```

## Problem 1
Write a function `digit_sum` that takes a positive integer and returns the sum of its digits.

<!-- grade_id: cell-38e8ef1defadda2c, points: 0, grade: false, solution: true -->
```python
def digit_sum(num):
    digits = [int(dig) for dig in str(num)]
    sum_dig = sum(digits)
    
    return sum_dig
```

<!-- grade_id: cell-5d5cfac174731d2c, points: 1, grade: true, solution: false -->
```python
assert_equal(digit_sum(0), 0)
assert_equal(digit_sum(2020), 4)
assert_equal(digit_sum(1234), 10)
```

## Problem 2
Write a function `count_vowels` that returns the number of vowels in an input string.

<!-- grade_id: cell-4d0b4cdd7f5064e0, points: 0, grade: false, solution: true -->
```python
def count_vowels(string):
    vowels = ['a','e','i','o','u']
    
    word_vowels = [v for v in string if v in vowels]
    
    return len(word_vowels)
```

<!-- grade_id: cell-1fa31947d992477b, points: 1, grade: true, solution: false -->
```python
assert_equal(count_vowels('foo bar baz'), 4)
assert_equal(count_vowels('fbrbz'), 0)
```

## Problem 3
Two words of the same length "interlock" if taking alternating letters from each forms a new word. For example, "shoe" and "cold" interlock to form "schooled".

Write a function that checks whether a pair of words interlock based on a word list.

<!-- grade_id: cell-d595349223cd3695, points: 0, grade: false, solution: true -->
```python
def is_interlock(word_list, word1, word2):
    '''
    Check if word1 and word2 interlocks based on word_list

    Two words "interlock" if taking alternating letters from each forms a new 
    word. For example, "shoe" and "cold" interlock to form "schooled".
    
    
    Parameters
    ----------
    word_list : list
        List of valid words
    word1 : string
        First word to check
    word2 : string
        Other word to check
        
        
    Returns
    -------
    interlockness : bool
        True if `word1` and `word2` interlock
    '''
    li_combined_1 = [i+v for i,v in zip(word1,word2)]
    li_combined_2 = [i+v for i,v in zip(word2,word1)]
    new_word_1 = ''
    new_word_2 = ''
    for i,j in zip(li_combined_1,li_combined_2):
        new_word_1 = new_word_1 + i
        new_word_2 = new_word_2 + j
    
    if new_word_1 in word_list or new_word_2 in word_list:
        return True
    else:
        return False
    
```

<!-- grade_id: cell-951f0bf1ac155acd, points: 1, grade: true, solution: false -->
```python
assert_equal(
    is_interlock(['school', 'schooled', 'shoe', 'cold'], 'shoe', 'cold'), 
    True)
assert_equal(
    is_interlock(['school', 'schooled', 'shoe', 'cold'], 'cold', 'shoe'), 
    True)
assert_equal(
    is_interlock(['school', 'shoe', 'cold'], 'shoe', 'cold'), 
    False)
assert_equal(is_interlock(['school'], 'shoe', 'cold'), False)
```

## Problem 4

Write a function `count_types` that accepts a string `a_string`, and returns the number of `lowercase`, `uppercase`, `numeric`, `punctuation` and `whitespace` characters in `a_string` as items in a dictionary.

<!-- grade_id: cell-7ce601f79cc30766, points: 0, grade: false, solution: true -->
```python
def count_types(a_string):
    punc_set = """!"#$%&'()*+,-./:;<=>?@\^_`{|}~"""

    count_space = count_lower = count_upper = count_num = count_punc = 0

    for char in a_string:

        if char.islower() is True:
            count_lower = count_lower + 1

        elif char.isupper() is True:
            count_upper = count_upper + 1

        elif char.isnumeric() is True:
            count_num = count_num + 1

        elif char in punc_set:
            count_punc = count_punc + 1

        else:
            count_space = count_space + 1

    d = {'lowercase': count_lower,
        'uppercase': count_upper,
        'numeric': count_num,
        'punctuation': count_punc,
        'whitespace': count_space}
    
    return d
```

<!-- grade_id: cell-9f6487557e54f1a9, points: 1, grade: true, solution: false -->
```python
assert_equal(
    count_types('abcABC123!!! '),
    {'lowercase': 3,
     'uppercase': 3,
     'numeric': 3,
     'punctuation': 3,
     'whitespace': 1}
)
assert_equal(
    count_types('abc123!'),
    {'lowercase': 3,
     'uppercase': 0,
     'numeric': 3,
     'punctuation': 1,
     'whitespace': 0}
)
assert_equal(
    count_types("""Mary had a little lamb
Little lamb, little lamb
Mary had a little lamb
It's fleece was white as snow"""),
    {'lowercase': 75,
     'uppercase': 4,
     'numeric': 0,
     'punctuation': 2,
     'whitespace': 19}
)
```

## Problem 5

Write a function `matmul` that accepts two matrices as list of lists and returns their matrix product as a list of lists. Assume the matrices are nonempty and have compatible dimensions.

<!-- grade_id: cell-5572a8301703f175, points: 0, grade: false, solution: true -->
```python
def matmul(list1,list2):
    
    product = [[sum(a*b for a,b in zip(list1_row, list2_col)) for list2_col in zip(*list2)] for list1_row in  list1]
    return product

```

<!-- grade_id: cell-902b3b435df6c97c, points: 1, grade: true, solution: false -->
```python
assert_equal(
    matmul(
        [[1, 2],
         [3, 4]],
        [[5, 6],
         [7, 8]]
    ),
    [[19, 22],
     [43, 50]]
)
assert_equal(
    matmul(
        [[1, 2, 7],
         [3, 4, 8]],
        [[5, 6],
         [7, 8],
         [9, 10]]
    ),
    [[82, 92], 
     [115, 130]]
)
assert_equal(
    matmul(
        [[1, 2],
         [3, 4],
         [7, 8]],
        [[5, 6, 9],
         [7, 8, 10]]
    ),
    [[19, 22, 29], 
     [43, 50, 67], 
     [91, 106, 143]]
)
```

## Problem 6
One classic method for composing secret messages is called a square code.  The spaces are removed from the english text and the characters are written into a square (or rectangle).  For example, the sentence "If man was meant to stay on the ground god would have given us roots" is 54 characters long, so it is written into a rectangle with 7 rows and 8 columns. 

    ifmanwas
    meanttos        
    tayonthe
    groundgo
    dwouldha
    vegivenu
    sroots
                
The coded message is obtained by reading down the columns going left to right.   For example, the message above is coded as:

    imtgdvs fearwer mayoogo anouuio ntnnlvt wttddes aohghn sseoau

Write a function `encode` that takes in a a string message with no spaces between the words and returns the encode message.

<!-- grade_id: cell-673ab947a5fc9c66, points: 0, grade: false, solution: true -->
```python
def encode(text):
    x = [i.lower() for i in text if i.isalpha() or i.isdigit()]
    col = int((len(x)**(0.5))) + 1
    if col * (col - 1) >= len(x):
        row = col - 1
    else:
        row = col
    
    res = []
    for i in range(col):
        for j in range(row):
            try:
                res.append(x[i + j * col])
            except:
                res.append('')
    for i in range (col - 1,0,-1):
        res.insert(i * row, ' ')
    return ''.join(res)
```

<!-- grade_id: cell-8498f3626841532e, points: 1, grade: true, solution: false -->
```python
assert_equal(encode('haveaniceday'), 'hae and via ecy')
assert_equal(encode('feedthedog'), 'fto ehg ee dd')
assert_equal(encode('chillout'), 'clu hlt io')
```

## Problem 7
We are given strings containing brackets of four types - round (), square [], curly {} and angle <> ones. Create a function `check_brackets` that checks whether brackets are in correct sequence i.e., any opening bracket should have closing bracket of the same type somewhere further by the string, and bracket pairs should not overlap, though they could be nested:

    (a+[b*c] - {d/3})  - here square and curly brackets are nested in the round ones

    (a+[b*c) - 17]     - here square brackets overlap with round ones which does not make sense

<!-- grade_id: cell-c457791d0d5db697, points: 0, grade: false, solution: true -->
```python
def check_brackets(str_with_brackets):
    """Check whether str_with_bracks is bracketed correctly
    
    Parameters
    ----------
    str_with_brackets : str
        String with brackets that are possibly nested
    
    Returns
    -------
    is_correct : bool
        `True` if `str_with_brackets` is bracketed correctly, `False` 
        otherwise
    """
    open_brackets = ["[","{","("]
    close_brackets = ["]","}",")"]
    dump = []
    for i in str_with_brackets:
        if i in open_brackets:
            dump.append(i)
        elif i in close_brackets:
            pos = close_brackets.index(i)
            if ((len(dump) > 0) and 
                (open_brackets[pos] == dump[len(dump)-1])):
                dump.pop()
            else:
                return False
    if len(dump) == 0:
        return True
    else:
        return False
```

<!-- grade_id: cell-6c9fa78273cf19d3, points: 1, grade: true, solution: false -->
```python
assert_equal(check_brackets('(a+[b*c] - {d/3})'), True)
assert_equal(check_brackets('(a+[b*c) - 17]'), False)
```

## Problem 8
Write a function called `nested_sum` that takes a list of lists of integers and adds up the elements from all of the nested lists. For example:

    >>> t = [[1, 2], [3], [4, 5, 6]]
    >>> nested_sum(t)
    21

<!-- grade_id: cell-173f8231b323732b, points: 0, grade: false, solution: true -->
```python
def nested_sum(list_of_lists):
    # YOUR CODE HERE
    total = 0
    for element in list_of_lists:
        if (type(element) == type([])):
            total = total + nested_sum(element)
        else:
            total = total + element
    return total
```

<!-- grade_id: cell-eae30f88ad6a0851, points: 1, grade: true, solution: false -->
```python
assert_equal(nested_sum([[1, 2], [3], [4, 5, 6]]), 21)
```

## Problem 9

Write a function `count_people` that accepts a multiline string of log entries. Each log entry consists of the string `IN` or `OUT` corresponding to whether the entry is the number of people who went in or out of the building, respectively, and the number of people this entry applies to, separate by the tab character (`\t`). `count_people` returns the number of people inside the building based on the log.

<!-- grade_id: cell-03dea1b273a41cb5, points: 0, grade: false, solution: true -->
```python
def count_people(log):
    count = 0
    text = log.split()
    pairs = []
    for i in range (0,len(text),2):
        pairs += [(text[i], text[i+1])]
    
    for i in pairs:
        if i[0] == 'IN':
            pos = pairs.index(i)
    
    for i in pairs[pos::]:
        if i[0] == 'IN':
            count += int(i[1])
        else:
            count -= int(i[1])
    return count
```

<!-- grade_id: cell-95ad0d9628941109, points: 1, grade: true, solution: false -->
```python
assert_equal(
    count_people(
        'IN\t1\n'
        'OUT\t1\n'
        'IN\t1\n'
        'IN\t1\n'
        'IN\t1\n'
        'IN\t1\n'
        'OUT\t2\n'
    ),
    2
)
```

## Problem 10
Write a function `next_word` that takes a string of words and returns the most likely next word for all words or for a single word.

<!-- grade_id: cell-ad39c27c5c5f2d1e, points: 0, grade: false, solution: true -->
```python
from collections import Counter
def next_word(text, word=None):
    """
    Return the most likely next word in text
    
    A word is defined as a sequence of all non-whitespace characters between
    whitespaces. Words are case-insensitive.
    
    Parameters
    ----------
    text : string
        Text to train at.
    word : string or `None`
        Find the most likely next word of `word` or likely next word of all
        words if `None`. 
        
    
    Returns
    -------
    next_word : tuple or list of tuple
        If `word` is a string then return the most likely next word of `word`
        as a tuple of `(word, most_likely_next_word)`. If `word` is not found
        in `text`, `most_likely_next_word` is an empty string. If `word` is
        `None` then return the list of `(word, most_likely_next_word)` for all
        words in `text`. If there is more than one most likely next word, pick
        the first word based on alphabetical (lexicographic) order.
    """
    words = []
    nxt = []
    dummy = []
    temp = []
    if isinstance(word, str) and word in text:
        words = text.split()
        n = len(words)

        for i in range(len(words)):
            if words[i] == word:
                try:
                    nxt.append((words[i], words[i + 1]))
                except:
                    continue
        return nxt[0]
    elif word is None:
        words = text.split()
        n = len(words)
        for i in range(len(words)-1):
            try:
                nxt.append((words[i], words[i + 1]))
            except:
                continue
        nxt = sorted(nxt, key = lambda x: (x[0],x[1]))
        count = dict(Counter(nxt))
        sorted_keys = sorted(count, key=count.get, reverse = True)
        sorted_dict = {}

        for i in sorted_keys:
            sorted_dict[i] = count[i]
        
        for i in sorted_dict:
            if i[0] in dummy:
                continue
            else:
                dummy.append(i[0])
                temp.append(i)
        nxt = temp
        return nxt
```

<!-- grade_id: cell-8721e2955dbdaef8, points: 1, grade: true, solution: false -->
```python
assert_equal(
    set(next_word('a b c d a b c d e f a b c a c a c d e a')),
    set([('a', 'b'), ('b', 'c'), ('c', 'd'), 
         ('d', 'e'), ('e', 'a'), ('f', 'a')]))
assert_equal(
    next_word('a b c d a b c d e f a b c a c a c d e a', 'a'), 
    ('a', 'b'))
assert_equal(
    next_word('a b c d a b c d e f a b c a c a c d e a', 'b'), 
    ('b', 'c'))
assert_equal(
    next_word('a b c d a b c d e f a b c a c a c d e a', 'c'), 
    ('c', 'd'))
```

```python

```