```python
from numpy.testing import assert_equal
```

For all items, the answer should be a single-line Linux (bash) shell command. Concatenation using `&&` and `;` is not allowed.

## Problem 1

Show the $(n-20)$ to $(n - 16)$ (inclusive) lines, where $n$ is the last line in the file, of "The Adventures of Sherlock Holmes", which is the file `1/6/6/1661/1661.txt` under the `gutenberg` public data directory.

<!-- grade_id: cell-86cc8b6f7b819256, points: 0, grade: false, solution: true -->
```python
%%bash --out output
# lines in file = 13053
cat /mnt/data/public/gutenberg/1/6/6/1661/1661.txt | tail -n 21 | head -n 5

```

<!-- grade_id: prob1-test, points: 1, grade: true, solution: false -->
```python
lines = output.split('\r\n')
assert_equal(lines[0], '')
```

## Problem 2
How many items belong to the category `Drugs/Cannabis/Weed` in the Agora dark web marketplace? Use the data found in the public data directory.

<!-- grade_id: cell-e28d4638d1afa2de, points: 0, grade: false, solution: true -->
```python
%%bash --out output
cat /mnt/data/public/agora/Agora.csv | grep -o Drugs/Cannabis/Weed | wc -l
```

<!-- grade_id: prob2-test, points: 1, grade: true, solution: false -->
```python
assert_equal(output.strip(), '21727')
```

## Problem 3
Create the directory structure shown below from the current directory:

    a/b/c
    d/e/f
    j/k/l

<!-- grade_id: cell-e9663008a48ff85a, points: 0, grade: false, solution: true -->
```python
%%bash --out output
rm -rf a
rm -rf d
rm -rf j

mkdir -p a/b/c d/e/f j/k/l
```

<!-- grade_id: prob3-test, points: 1, grade: true, solution: false -->
```python
import os
assert_equal(os.path.exists('a/b/c'), True)
assert_equal(os.path.exists('d/e/f'), True)
assert_equal(os.path.exists('j/k/l'), True)
```

## Problem 4
Display only the details of the largest file inside the Amazon Reviews public data directory. The file size should be in human-readable sizes (e.g., 1K 234M 2G).

<!-- grade_id: cell-f3b369e81d39f74a, points: 0, grade: false, solution: true -->
```python
%%bash --out output
ls -Slh /mnt/data/public/amazon-reviews/ | head -n 2 | tail -n 1
```

<!-- grade_id: prob4-test, points: 1, grade: true, solution: false -->
```python
lines = output.strip().split('\n')
assert_equal(
    lines[0], 
    '-rw-rw-rw- 1 root root 2.6G Aug 16  2018 '
    'amazon_reviews_us_Books_v1_00.tsv.gz')
assert_equal(len(lines), 1)
```

## Problem 5
List down the sorted filepaths of h5 files with track ID starting with `A` in the Million Song dataset.

<!-- grade_id: cell-677401e3862c4a05, points: 0, grade: false, solution: true -->
```python
%%bash --out output
find /mnt/data/public/millionsong/*A -type f -name "*.h5" | sort 
```

<!-- grade_id: prob5-test, points: 1, grade: true, solution: false -->
```python
lines = output.strip().split('\n')
assert_equal(len(lines), 39100)
assert_equal(
    lines[0], 
    '/mnt/data/public/millionsong/A/A/A/TRAAAAK128F9318786.h5')
```

## Problem 6
Show the file size in bytes and absolute filepath of files under the `census` public data directory which are more than 1000000 bytes. Sort them by increasing file size.

<!-- grade_id: cell-60ce3d081607f4c8, points: 0, grade: false, solution: true -->
```python
%%bash --out output
find /mnt/data/public/census/ -size +1000000c -printf %s" " -printf %p\\n | \
sort
```

<!-- grade_id: prob6-test, points: 1, grade: true, solution: false -->
```python
lines = output.strip().split('\n')
assert_equal(
    lines[0], 
    '1004032 /mnt/data/public/census/Northern Samar_Statistical Tables.xls')
assert_equal(len(lines), 25)
```

## Problem 7
List down all the mentions files for 14 Feb 2019 between 8am to 9am (exclusive) in the file `masterfilelist.txt` under the GDELT v2 public data directory.

<!-- grade_id: cell-a1fb49df85aeb796, points: 0, grade: false, solution: true -->
```python
%%bash --out output
cat /mnt/data/public/gdeltv2/masterfilelist.txt | grep -i 201902148*| \
grep -i mentions| grep -in mentions| head -n 36 | tail -n 4 | \
cut -d ' ' -f 3
```

<!-- grade_id: prob7-test, points: 1, grade: true, solution: false -->
```python
lines = output.strip().split('\n')
assert_equal(
    lines[0],
    'http://data.gdeltproject.org/gdeltv2/20190214080000.mentions.CSV.zip')
```

## Problem 8
Create a clone copy of the `mnist` public data directory to the current directory. Being a clone copy, it must preserve all attributes (mode, ownership, timestamps), links and directory structure.

Note: the answer cell is a markdown cell to prevent you from actually copying the directory. This problem will be manually checked.

<!-- grade_id: prob8, points: 1, grade: true, solution: true -->
cp -rp /mnt/data/public/mnist . 

## Problem 9
How many unique tags are there in `/mnt/data/public/movielens/20m/ml-20m/tags.csv`?

<!-- grade_id: cell-3fe975bd03ab9a17, points: 0, grade: false, solution: true -->
```python
%%bash --out output
cat /mnt/data/public/movielens/20m/ml-20m/tags.csv | \
cut -d ',' -f 3| sort | uniq | wc -l
```

<!-- grade_id: prob9-test, points: 1, grade: true, solution: false -->
```python
assert_equal(output.strip(), '38641')
```

## Problem 10
Store the list of all book titles that contain the word `pandemic` (case-insensitive) in `BX-Books.csv` under the book-crossing public data directory to the file `pandemic-books.txt`.

```python
!rm -f pandemic-books.txt
```

<!-- grade_id: cell-7717d22f83c79ba0, points: 0, grade: false, solution: true -->
```python
%%bash
grep -i pandemic /mnt/data/public/book-crossing/BX-Books.csv | cut -d ';' \
-f 2 > pandemic-books.txt
```

<!-- grade_id: prob10-test, points: 1, grade: true, solution: false -->
```python
res = open('pandemic-books.txt').readlines()
assert_equal(len(res), 6)
assert_equal(
    res[0], 
    '"Flu: The Story of the Great Influenza Pandemic of 1918 and the Search for the Virus That Caused It"\n')
```