```python
from pyspark import SparkContext
from pyspark.sql import SparkSession
from numpy.testing import assert_equal
```

```python
import os
os.environ['AWS_PROFILE'] = 'default'
```

```python
from pyspark import SparkConf
```

```python
conf = (SparkConf()
           .setMaster('local[*]')
           .set('spark.jars.packages',
                'com.amazonaws:aws-java-sdk:1.11.563,org.apache.hadoop:hadoop-aws:3.2.2')
           .set('spark.executor.extraJavaOptions',
                '-Dcom.amazonaws.services.s3.enableV4=true')
           .set('spark.driver.extraJavaOptions',
                '-Dcom.amazonaws.services.s3.enableV4=true')
           .set('fs.s3a.aws.credentials.provider',
                'com.amazonaws.auth.profile.ProfileCredentialsProvider,'
                'org.apache.hadoop.fs.s3a.AnonymousAWSCredentialsProvider'))
sc = SparkContext(conf=conf)
```

Output:
```
Warning: Ignoring non-Spark config property: fs.s3a.aws.credentials.provider
WARNING: An illegal reflective access operation has occurred
WARNING: Illegal reflective access by org.apache.spark.unsafe.Platform (file:/usr/local/spark-3.1.2-bin-hadoop3.2/jars/spark-unsafe_2.12-3.1.2.jar) to constructor java.nio.DirectByteBuffer(long,int)
WARNING: Please consider reporting this to the maintainers of org.apache.spark.unsafe.Platform
WARNING: Use --illegal-access=warn to enable warnings of further illegal reflective access operations
WARNING: All illegal access operations will be denied in a future release
:: loading settings :: url = jar:file:/usr/local/spark-3.1.2-bin-hadoop3.2/jars/ivy-2.4.0.jar!/org/apache/ivy/core/settings/ivysettings.xml
Ivy Default Cache set to: /home/msds2022/jgacal/.ivy2/cache
The jars for the packages stored in: /home/msds2022/jgacal/.ivy2/jars
com.amazonaws#aws-java-sdk added as a dependency
org.apache.hadoop#hadoop-aws added as a dependency
:: resolving dependenc
```

```python
spark = SparkSession(sc)
```

```python

```

```python
!aws s3 ls s3://helpful-sentences-from-reviews/
```

Output:
```
2021-05-05 14:26:24          0 README.md
2021-05-05 14:26:23          0 index.html
2021-04-18 13:08:05     525322 test.json
2021-04-18 13:08:23    5258801 train.json
```

```python

```

# Problem 2

For this problem, you will be working with the [Helpful Sentences from Reviews](https://registry.opendata.aws/helpful-sentences-from-reviews/) training set.

You may **only** use Apache Spark and the Python Standard Library. You **cannot** use numpy, scipy, pandas or scikit-learn. Do **not** print or display large amount of results. You will get deductions if you make the browser unresponsive.

## Problem 2a [10 pts]

Create a function `count_ratings` that reads `train.json` and returns a Spark DataFrame with columns `rating`, which is the `helpful` value rounded to one decimal place, and `count`, the number of sentences with that rating. Sort by `rating`.

**correction: business rounding dapat**

```python
import pyspark.sql.functions as f
```

```python
from decimal import Decimal



# df = spark.read.json('s3a://helpful-sentences-from-reviews/train.json')
# (df
#  .withColumn('rating', f.udf(lambda x:float(Decimal(str(x)).quantize(Decimal('1.0'))))(f.col('helpful')))
#  .groupBy('rating')
#  .count()
#  .orderBy('rating')
#  .limit(10)
 
#  .toPandas()
# )
```

```python
# df = spark.read.json()
```

```python
from pyspark.sql.functions import udf
@udf
def b_round(x):
    import decimal as d
    return (d.Decimal(str(x)).quantize(d.Decimal('1.0')))
```

<!-- grade_id: cell-5098bd4de40d844d, points: 0, grade: false, solution: true -->
```python
def count_ratings():
    from pyspark.sql.types import FloatType
    return (spark.read.json('s3a://helpful-sentences-from-reviews/train.json')
            
 .withColumn('rating', f.round(b_round(f.col('helpful')), 1))
 .groupBy('rating')
 .count()
#  .withColumn('rating', f.udf(lambda x: x*1.0)(f.col('rating').cast('float')))
 .orderBy('rating')
           )
```

<!-- grade_id: cell-2095fd1c2996b071, points: 10, grade: true, solution: false -->
```python
pdf_ratings_count = count_ratings().toPandas()
assert_equal(len(pdf_ratings_count), 21)
assert_equal(pdf_ratings_count.columns.tolist(), ['rating', 'count'])
assert_equal(
    pdf_ratings_count.iloc[:10].to_numpy().tolist(),
    [[0.0, 37.0],
     [0.1, 91.0],
     [0.2, 188.0],
     [0.3, 272.0],
     [0.4, 452.0],
     [0.5, 615.0],
     [0.6, 797.0],
     [0.7, 970.0],
     [0.8, 1122.0],
     [0.9, 1376.0]]
)
```

Output:
```
22/03/07 16:21:45 WARN MetricsConfig: Cannot locate configuration: tried hadoop-metrics2-s3a-file-system.properties,hadoop-metrics2.properties
```

## Problem 2b [15 pts]

Cluster the sentences by their `sentence` text using k-means. You may preprocess it however you want. Select the optimal number of clusters using silhouette coefficient. Make sure that the steps that you are doing are sufficiently documented, and the final model and number of clusters are explicit.

```python
df = spark.read.json('s3a://helpful-sentences-from-reviews/train.json')
df.limit(4).toPandas()
```

Output:
```
asin  helpful                                     main_image_url  \
0  B000AO3L84     1.70  http://ecx.images-amazon.com/images/I/41XAEKR9...   
1  B001SEQPGK     1.30  http://ecx.images-amazon.com/images/I/71KLvmtc...   
2  0553386697     1.90  http://ecx.images-amazon.com/images/I/81HdbmkR...   
3  B006SUWZH2     0.25  http://ecx.images-amazon.com/images/I/61A2WQOL...   

                                       product_title  \
0  Canon 430EX Speedlite Flash for Canon EOS SLR ...   
1  Sony Cyber-shot DSC-W290 12 MP Digital Camera ...   
2  The Whole-Brain Child: 12 Revolutionary Strate...   
3       Memorex Portable CD Boombox with AM FM Radio   

                                            sentence  
0                      this flash is a superb value.  
1                The pictures were not sharp at all.  
2                  A very good resource for parents.  
3  We have it in a child's room, and will be swit...
```

```python
df_sen = (df
 .select('sentence')
 .dropna()
)

df_sen.limit(4).toPandas()
```

Output:
```
sentence
0                      this flash is a superb value.
1                The pictures were not sharp at all.
2                  A very good resource for parents.
3  We have it in a child's room, and will be swit...
```

```python
from pyspark.ml.clustering import KMeans
from pyspark.ml.evaluation import ClusteringEvaluator
from pyspark.ml.feature import HashingTF, IDF, Tokenizer
from pyspark.ml import Pipeline

```

perform tfidf

```python
tokenizer = Tokenizer(inputCol="sentence", outputCol="words")
hashingTF = HashingTF(inputCol=tokenizer.getOutputCol(),outputCol="rawFeatures", numFeatures=30)
idf = IDF(inputCol=hashingTF.getOutputCol(), outputCol="features",minDocFreq=5)
pipeline = Pipeline(stages=[tokenizer, hashingTF, idf, KMeans(k=8)])

transformer = pipeline.fit(df_sen)
df_transformed = transformer.transform(df_sen)

df_transformed.limit(5).toPandas()
```

Output:
```
22/03/07 16:22:42 WARN BLAS: Failed to load implementation from: com.github.fommil.netlib.NativeSystemBLAS
22/03/07 16:22:42 WARN BLAS: Failed to load implementation from: com.github.fommil.netlib.NativeRefBLAS
                                                                                                                            sentence  \
0                      this flash is a superb value.   
1                The pictures were not sharp at all.   
2                  A very good resource for parents.   
3  We have it in a child's room, and will be swit...   
4  Again the makers are too lazy to bring in the ...   

                                               words  \
0               [this, flash, is, a, superb, value.]   
1        [the, pictures, were, not, sharp, at, all.]   
2           [a, very, good, resource, for, parents.]   
3  [we, have, it, in, a, child's, room,, and, wil...   
4  [again, the, makers, are, too, lazy, to, bring...   

                                  
```

```python
evaluator = ClusteringEvaluator()
```

perform CV

```python
from pyspark.ml.tuning import CrossValidator, ParamGridBuilder

paramGrid = (ParamGridBuilder()
                .addGrid(pipeline.getStages()[-1].k, [2,3, 4,5,6,7, 8]) #add grid here specificies the variable you want to change
                .build())

cv = CrossValidator(estimator=pipeline,
                    estimatorParamMaps=paramGrid,
                    evaluator=evaluator)

cv_trained=cv.fit(df_sen)
```

```python
cv_trained.bestModel.stages[-1]
```

Output:
```
KMeansModel: uid=KMeans_4a94fc0b4758, k=2, distanceMeasure=euclidean, numFeatures=30
```

```python
cv_trained.bestModel.stages[-1].explainParam('k')
```

Output:
```
'k: The number of clusters to create. Must be > 1. (default: 2, current: 2)'
```

best k is 2. Metric used is silhouette coeff by default.