```python
import numpy as np
from numpy.testing import (assert_equal, assert_raises, assert_array_equal, 
                           assert_almost_equal, assert_allclose)
```

## Problem 1
Create a function `div` that accepts `a` and `b` as parameters and returns `a / b`. If `b` is zero then return `np.nan`. Other exceptions should remain unhandled.

<!-- grade_id: cell-5c3bb762f90f6978, points: 0, grade: false, solution: true -->
```python
def div(a, b):
    """
    Compute a divided by b
    
    Accepts a and b as parameters and returns a / b. If b is zero then return
    np.nan. Other exceptions remain unhandled.
    
    Parameters
    ----------
    a : int
        numerator
    b : int
        denominator
    
    Returns
    -------
    float
        quotient
    """
    try:
        a/b
    except:
        return np.nan
    else:
        return a/b
```

<!-- grade_id: cell-1f0fc02049cc7b3e, points: 1, grade: true, solution: false -->
```python
assert_almost_equal(div(1, 2), 0.5)
assert_equal(div(10, 0), np.nan)
```

## Problem 2
Create a function `gen_array` that returns a 10$\times$10 `numpy` `ndarray` from numbers 1 to 100 with values divisible by 3 replaced by 0:

     array([1  2 0  … 0  10],
           [11 …     19  20],
            …
           [91 …      0 100])

<!-- grade_id: cell-2ff687ff756458eb, points: 0, grade: false, solution: true -->
```python
def gen_array():
    """
    Returns a 10 × 10 numpy ndarray 
    
    ndarray contains numbers 1 to 100 with values divisible by 3 replaced by 0
    
    Returns
    -------
    arr2 : ndarray
        10 × 10 numpy ndarray containing the values
    """
    li2 = [i if i%3 != 0 else 0 for i in range(1,101)]
    arr2 = np.array(li2).reshape(10,10)

    return arr2
```

```python

```

Output:
```
array([[  1,   2,   0,   4,   5,   0,   7,   8,   0,  10],
       [ 11,   0,  13,  14,   0,  16,  17,   0,  19,  20],
       [  0,  22,  23,   0,  25,  26,   0,  28,  29,   0],
       [ 31,  32,   0,  34,  35,   0,  37,  38,   0,  40],
       [ 41,   0,  43,  44,   0,  46,  47,   0,  49,  50],
       [  0,  52,  53,   0,  55,  56,   0,  58,  59,   0],
       [ 61,  62,   0,  64,  65,   0,  67,  68,   0,  70],
       [ 71,   0,  73,  74,   0,  76,  77,   0,  79,  80],
       [  0,  82,  83,   0,  85,  86,   0,  88,  89,   0],
       [ 91,  92,   0,  94,  95,   0,  97,  98,   0, 100]])
```

<!-- grade_id: cell-178ae4e110662a47, points: 1, grade: true, solution: false -->
```python
a = gen_array()
assert_equal(isinstance(a, np.ndarray), True)
assert_equal(a.dtype, int)
```

## Problem 3
Create a function `dot` that takes in two 1$\times$9 `numpy` `ndarray`s, converts the arrays into 3$\times$3 (e.g., `[1, 2, 3, 4, 5, 6, 7, 8, 9]` becomes `[[1, 2, 3], [4, 5, 6], [7, 8, 9]]` ) then returns the dot product of the arrays. The function should raise a `ValueError` if the inputs are not 1$\times$9 `ndarray`s.

<!-- grade_id: cell-c5cd80f0f11ad2a2, points: 0, grade: false, solution: true -->
```python
def dot(arr1, arr2):
    """
    Convert 2 1x9 input arrays into 3x3 arrays and return their dot product 
    
    Raises a ValueError if the inputs are not 1x9 ndarrays
    
    Parameters
    ----------
    arr1, arr2: ndarray
        ndarrays
    
    Returns
    -------
    ndarray
        resulting 3x3 ndarray after performing dot product on arr1 and arr2
    """
    if (type(arr1) is np.ndarray) and (type(arr2) is np.ndarray):
        if (arr1.shape == (1,9) and arr2.shape == (1,9)):
            arr1 = arr1.reshape((3,3))
            arr2 = arr2.reshape((3,3))
        else:
            raise ValueError
    else:
        raise ValueError
    
    return np.dot(arr1,arr2)

```

```python

```

<!-- grade_id: cell-421df91b38cd545f, points: 1, grade: true, solution: false -->
```python
assert_raises(ValueError, lambda: dot(range(9), np.arange(9)))
```

## Problem 4
Create a function `mult3d` that takes in a 3D `ndarray` `a` and a square `ndarray` `b` and returns an `ndarray` `z` defined by $z_{ijk} = a_{ijk}b_{ij}$.

<!-- grade_id: cell-88c1cae07be633da, points: 0, grade: false, solution: true -->
```python
def mult3d(a, b):
    """
    Create ndarray z, where z(i,j,k) is the product of a(i,j,k) and b(i,j,k)
    
    
    Parameters
    ----------
    a : ndarray
        3D ndarray
    b : ndarray
        square ndarray
    
    Returns
    -------
    ndarray
        resulting 3D ndarray
    """
    b = b.reshape(b.shape + (1,))
    return b*a
```

<!-- grade_id: cell-db0692031e129076, points: 1, grade: true, solution: false -->
```python
a = np.ones((2, 2, 3))
b = np.array([[1, 2], [3, 4]])
z = mult3d(a, b)
assert_equal(isinstance(z, np.ndarray), True)
```

## Problem 5

Create a class `ABMWithStep` that is an `ABM` but returns the string `step` with the step number when `status` is called. 

```python
class ABM:
    def __init__(self):
        self.timestep = 0
        
    def step(self):
        self.timestep += 1
        
    def status(self):
        return 'step'

def step_model(model, steps):
    statuses = []
    for _ in range(steps):
        statuses.append(model.status())
        model.step()
    return statuses
```

<!-- grade_id: cell-11bce336caf0f1b4, points: 0, grade: false, solution: true -->
```python
class ABMWithStep(ABM):
    """
    Child of class ABM
    """
    def status(self):
        """
        return 'step x', where x is the timestep of string type

        Returns
        -------
        str
            step number
        """
        return 'step ' + str(self.timestep)
```

<!-- grade_id: cell-95c2acb09668d30d, points: 1, grade: true, solution: false -->
```python
abm = ABM()
assert_equal(
    step_model(abm, 10), 
    ['step', 'step', 'step', 'step', 'step', 
     'step', 'step', 'step', 'step', 'step'])
abm_ws = ABMWithStep()
assert_equal(isinstance(abm_ws, ABM), True)
assert_equal(
    step_model(abm_ws, 10), 
    ['step 0', 'step 1', 'step 2', 'step 3', 'step 4', 
     'step 5', 'step 6', 'step 7', 'step 8', 'step 9'])
```

## Problem 6

Create a class `FlightTracker` that is a `Tracker` but also accepts `height` as a required parameter on creation. It should also have a `get_height` method that returns the current `height`.

```python
class Tracker:
    def __init__(self, lon, lat):
        self.lon = lon
        self.lat = lat
        
    def get_position(self):
        return (self.lon, self.lat)
```

<!-- grade_id: cell-1a7ef50fce72f4b8, points: 0, grade: false, solution: true -->
```python
class FlightTracker(Tracker):
    """
    Child of class Tracker
    """
    def __init__(self, lon, lat, height):
        """
        Initialize longitute, latitude, and height values
        
        Parameters
        ----------
        lon : int
            longitude value
        lat : int
            latitude value
        height : int
            height value
        """
        self.height = height
        self.lon = lon
        self.lat = lat
    def get_height(self):
        """
        Return height value


        Returns
        -------
        int
            height value
        """
        return self.height
```

<!-- grade_id: cell-26c5b4815bfa00a0, points: 1, grade: true, solution: false -->
```python
tracker = Tracker(120, 10)
assert_equal(tracker.get_position(), (120, 10))
flight_tracker = FlightTracker(-60, 30, 30000)
assert_equal(isinstance(flight_tracker, Tracker), True)
assert_equal(flight_tracker.get_position(), (-60, 30))
assert_equal(flight_tracker.get_height(), 30000)
```

## Problem 7

Create a class `Rectangle` that is a `Polygon` but accepts the `length` and `width` on initialization. Create another class `Square` that is a `Rectangle` but accepts a `side` on initialization. Do not overwrite the `compute_perimeter` method of `Rectangle`.

```python
class Polygon:
    def __init__(self, *sides):
        self.sides = sides
        
    def compute_perimeter(self):
        return sum(self.sides)
```

<!-- grade_id: cell-7fc2e1822ff30382, points: 0, grade: false, solution: true -->
```python
class Rectangle(Polygon):
    """
    Child of class Polygon
    """
    def __init__(self, length, width):
        """
        Initialize length and width values; compute sides value
        
        Parameters
        ----------
        length : int
            length value
            
        width : int
            width value
        """
        self.length = length
        self.width = width
        self.sides = (2*length, 2*width)
class Square(Rectangle):
    """
    Child of class Rectangle
    """
    def __init__(self, side):
        """
        Initialize side value; compute sides value
        
        Parameters
        ----------
        side : int
            side value
        """
        self.side = side
        self.sides = (side*4,)
```

<!-- grade_id: cell-828fbee742a9e69e, points: 1, grade: true, solution: false -->
```python
assert_equal(issubclass(Rectangle, Polygon), True)
assert_equal(issubclass(Square, Rectangle), True)
assert_equal(Rectangle.compute_perimeter, Polygon.compute_perimeter)
assert_equal(Rectangle(4, 5).compute_perimeter(), 18)
assert_equal(Square.compute_perimeter, Polygon.compute_perimeter)
assert_equal(Square(5).compute_perimeter(), 20)
```

## Problem 8

In the module `transition`, create a class `TransitionMatrix` that is initialized by a `numpy` `ndarray` of transition matrix. The initialization should raise a `TypeError` if the passed transition matrix is not a `numpy` `ndarray`, and a `ValueError` if any element is not between 0 and 1 (both inclusive). `TransitionMatrix`  has a property `probabilities` that is initialized to the transition matrix. `TransitionMatrix` should also have a parameterless method `step` that returns a `TransitionMatrix` that has the same transition matrix but with `probabilities` equal to the product of `probabilities` and the transition matrix. The original `probabilities` and transition matrix should not change when `step` is called.

```python

```

<!-- grade_id: cell-bbcb5d1ee6a1c2a9, points: 1, grade: true, solution: false -->
```python
from transition import TransitionMatrix
assert_raises(
    TypeError, 
    lambda: TransitionMatrix(
        [[0.1, 0.9], 
         [0.2, 0.8]]))
assert_raises(
    ValueError, 
    lambda: TransitionMatrix(
        np.array([[1, 9], 
                  [2, 8]])))

tm = TransitionMatrix(
    np.array([[0.1, 0.9], 
              [0.2, 0.8]]))
assert_allclose(
    tm.probabilities, 
    np.array([[0.1, 0.9], 
              [0.2, 0.8]]))

tm2 = tm.step()
assert_allclose(
    tm2.probabilities, 
    np.array([[0.01, 0.81],
              [0.04, 0.64]]))
assert_allclose(
    tm.probabilities, 
    np.array([[0.1, 0.9], 
              [0.2, 0.8]]))

assert_allclose(
    tm.step().step().probabilities, 
    np.array([[0.001, 0.729],
              [0.008, 0.512]]))
assert_allclose(
    tm2.probabilities, 
    np.array([[0.01, 0.81],
              [0.04, 0.64]]))
assert_allclose(
    tm.probabilities, 
    np.array([[0.1, 0.9], 
              [0.2, 0.8]]))

```

## Problem 9

Create a function `zero_crossings` that counts the number of times the elements of a sequence touch or cross the $x$-axis ($y = 0$).

<!-- grade_id: cell-af7273a23b45b2d6, points: 0, grade: false, solution: true -->
```python
def zero_crossings(readings):
    """
    Count the number of times an integer touches or crosses the 𝑥-axis (𝑦=0).
    
    Iterate through a sequence of integers and compute the count value of each
    iterated number.
    
    
    Parameters
    ----------
    readings : list
        list of integers
            
    Returns
    -------
    count : int
        
    """
    count = 0
    for i in range(len(readings)-1):
        a = readings[i]
        b = readings[i+1]
        if (a != b) and not ((a<0 and b<0) or (a>0 and b>0)) :
            if a + b <= abs(a) + abs(b):
                count = count + 1
    return count
```

<!-- grade_id: cell-e413d168f7dd7288, points: 1, grade: true, solution: false -->
```python
assert_equal(zero_crossings([1, -10, 1, 5, 0, 0, 0, -1, -5, -6, 6]), 5)
```

```python
zero_crossings([0,0,1])
```

Output:
```
1
```

## Problem 10

Create a function `outer_sum` that accepts the one-dimensional sequences `x` and `y` and returns the two-dimensional list `z` defined by $z_{ij} = x_i + y_j$.

<!-- grade_id: cell-05a1eec4fbaf6f7f, points: 0, grade: false, solution: true -->
```python
def outer_sum(x, y):
    """
    Create array z, such that z(i,j) is the sum of x(i) and y(j)
    
    Accept one-dimensional sequences x and y and return 
    a two-dimensional list z
    
    
    Parameters
    ----------
    x : list
        list of integers
    y : list
        list of integers
    Returns
    -------
    z : list
        2D list
        
    """
    x = np.array(x)
    y = np.array(y)
    arr10 = x.reshape((len(x),1)) + y
    output10 = [list(i) for i in arr10]
    return output10
```

```python

```

<!-- grade_id: cell-9f722fc407ddbd90, points: 1, grade: true, solution: false -->
```python
os = outer_sum([1, 2, 3], [10, 20, 30, 40])
assert_equal(isinstance(os, list), True)
assert_equal(
    os, 
    [[11, 21, 31, 41], 
     [12, 22, 32, 42], 
     [13, 23, 33, 43]])
```