For this set of exercises, all answers should be the result of a single call to a `re` function.

```python

```

```python
import re
from numpy.testing import assert_equal
```

## Problem 1
Create a function `is_date` that uses python regular expressions to check whether the input string follows any of the following date format:

* yyyy-mm-dd
* yyyy/mm/dd
* mm-dd-yyyy
* mm/dd/yyyy
* dd-mm-yyyy
* dd/mm/yyyy

It should return `True` if it does and `False` otherwise. No need to check whether the input is a valid date.

```python
%load_ext pycodestyle_magic
%pycodestyle_on
```

Output:
```
The pycodestyle_magic extension is already loaded. To reload it, use:
  %reload_ext pycodestyle_magic
```

<!-- grade_id: cell-fa81bd9b31310679, points: 0, grade: false, solution: true -->
```python
def is_date(test_str):
    """Return True if string input is in correct format; False otherwise

    Parameters
    ----------
    test_str : str
        Input string

    Returns
    -------
    Bool
    """
    YMD = (r'^(\d{4})(-)((0[1-9])|(1[0-2]))(-)'
           '((0[1-9])|([1-2][0-9])|(3[0-1]))$')
    YMD2 = (r'^(\d{4})(/)((0[1-9])|(1[0-2]))(/)'
            '((0[1-9])|([1-2][0-9])|(3[0-1]))$')
    MDY = (r'((0[1-9])|(1[0-2]))(-)'
           '((0[1-9])|([1-2][0-9])|(3[0-1]))(-)'
           '(\d{4})$')
    MDY2 = (r'((0[1-9])|(1[0-2]))(/)'
            '((0[1-9])|([1-2][0-9])|(3[0-1]))(/)'
            '(\d{4})$')
    DMY = (r'((0[1-9])|([1-2][0-9])|(3[0-1]))(-)'
           '((0[1-9])|(1[0-2]))(-)'
           '(\d{4})$')
    DMY2 = (r'((0[1-9])|([1-2][0-9])|(3[0-1]))(/)'
            '((0[1-9])|(1[0-2]))(/)'
            '(\d{4})$')
    return (re.match('|'.join([YMD, YMD2, MDY, DMY, MDY2, DMY2]), test_str)
            is not None)
```

Output:
```
19:14: W605 invalid escape sequence '\d'
22:15: W605 invalid escape sequence '\d'
25:14: W605 invalid escape sequence '\d'
28:15: W605 invalid escape sequence '\d'
```

```python
# yyyy-mm-dd
# yyyy/mm/dd
# mm-dd-yyyy
# mm/dd/yyyy
# dd-mm-yyyy
# dd/mm/yyyy
def is_date(text):
    # yyyy-mm-dd
    re.findall(r'[1-9][0-9]{3}[-/](?:(?:0[1-9])|(?:1[0-2]))[-/](?:(?:[0-2][0-9])|(?:3[01]))', text)
    ymd = r'[1-9][0-9]{3}[-/](?:(?:0[1-9])|(?:1[0-2]))[-/](?:(?:[0-2][0-9])|(?:3[01]))$'
    # yyyy-dd-mm
    re.findall(r'[1-9][0-9]{3}[-/](?:(?:[0-2][0-9])|(?:3[01]))[-/](?:(?:0[1-9])|(?:1[0-2]))', text)
    ydm = r'[1-9][0-9]{3}[-/](?:(?:[0-2][0-9])|(?:3[01]))[-/](?:(?:0[1-9])|(?:1[0-2]))$'
    # dd-mm-yyyy
    re.findall(r'(?:(?:[0-2][0-9])|(?:3[01]))[-/](?:(?:0[1-9])|(?:1[0-2]))[-/][1-9][0-9]{3}', text)
    dmy = r'(?:(?:[0-2][0-9])|(?:3[01]))[-/](?:(?:0[1-9])|(?:1[0-2]))[-/][1-9][0-9]{3}$'
    # mm-dd-yyyy
    re.findall(r'(?:(?:0[1-9])|(?:1[0-2]))[-/](?:(?:[0-2][0-9])|(?:3[01]))[-/][1-9][0-9]{3}', text)
    mdy = r'(?:(?:0[1-9])|(?:1[0-2]))[-/](?:(?:[0-2][0-9])|(?:3[01]))[-/][1-9][0-9]{3}$'

    compiled = fr'(?:{ymd})|(?:{ydm})|(?:{dmy})|(?:{mdy})'

    return re.match(compiled, text) is not None
```

```python
re.findall(r'[1-9][0-9]{3}[-/](?:(?:0[1-9])|(?:1[0-2]))[-/](?:(?:[0-2][0-9])|(?:3[01]))', '5-18-2020')
```

Output:
```
[]
```

```python
is_date('5-18-2020')
```

Output:
```
[]
```

<!-- grade_id: cell-ab95141581705770, points: 1, grade: true, solution: false -->
```python
assert_equal(is_date('2020-05-18'), True)
assert_equal(is_date('2020/05/18'), True)
assert_equal(is_date('05-18-2020'), True)
assert_equal(is_date('05/18/2020'), True)
assert_equal(is_date('18-05-2020'), True)
assert_equal(is_date('18/05/2020'), True)
assert_equal(is_date('5-18-2020'), False)
assert_equal(is_date('5/18/2020'), False)
assert_equal(is_date('18-5-2020'), False)
assert_equal(is_date('18/5/2020'), False)
assert_equal(is_date('2020-054-18'), False)
assert_equal(is_date('2020-05-181'), False)
assert_equal(is_date('202-05-18'), False)
```

## Problem 2
Create a function `find_data` that returns all occurrences of `data set` or `dataset` in `text` as a list of strings in the same order as they appear in `text`.

<!-- grade_id: cell-488abc2ed1aa5177, points: 0, grade: false, solution: true -->
```python
def find_data(text):
    """
    Returns all instances of data set or dataset in text as a list of strings
    in the same order as they appear

    Parameters
    ----------
    text : str
        Input string

    Returns
    -------
    List
    """
    return [i.group() for i in re.finditer(r'\bdataset\b|\bdata\sset\b', text)]
```

```python
text = """
My dataset is bigger than your data set. Your data set is so tiny, ten of your
data sets can fit inside my dataset. Ten of your data sets inside one of my
data set! Your data set doesn't stand a chance.
"""
def find_data(text):
    return re.findall(r'\bdata\s?set\b', text)
```

<!-- grade_id: cell-56d01f65a5255b8c, points: 1, grade: true, solution: false -->
```python
text = """
My dataset is bigger than your data set. Your data set is so tiny, ten of your
data sets can fit inside my dataset. Ten of your data sets inside one of my
data set! Your data set doesn't stand a chance.
"""
assert_equal(
    find_data(text), 
    ['dataset', 'data set', 'data set', 'dataset', 'data set', 'data set'])
```

## Problem 3
Create a function `find_data` that returns all phrases in `text` that begin with the word `little` and end with the word `lamb` as a list of strings in the same order as they appear in `text`.

<!-- grade_id: cell-d320f306907bed75, points: 0, grade: false, solution: true -->
```python
def find_lamb(text):
    """
    Return all phrases in text that begin with the word little and end with
    the word lamb as a list of strings in the same order as they appear in
    text

    Parameters
    ----------
    text : str
        Input string

    Returns
    -------
    List
    """
    return re.findall(r'\blittle\b.*?\blamb\b', text)
```

```python
text = """
Mary had a little lamb, little lamb, little lamb
Mary had a little brown lamb, little brown lamb, little brown lamb
Whose wool is as brown as wood
"""
re.findall(r'\blittle\b.*\slamb\b', text)
```

Output:
```
['little lamb, little lamb, little lamb',
 'little brown lamb, little brown lamb, little brown lamb']
```

<!-- grade_id: cell-baa90ee7745a4128, points: 1, grade: true, solution: false -->
```python
text = """
Mary had a little lamb, little lamb, little lamb
Mary had a little brown lamb, little brown lamb, little brown lamb
Whose wool is as brown as wood
"""
assert_equal(
    find_lamb(text), 
    ['little lamb',
     'little lamb',
     'little lamb',
     'little brown lamb',
     'little brown lamb',
     'little brown lamb'])
```

Output:
```
7:21: W291 trailing whitespace
```

## Problem 4
Create a function `repeat_alternate` that returns a string where every other word of `text` is repeated.

<!-- grade_id: cell-0292708fb072d973, points: 0, grade: false, solution: true -->
```python
def repeat_alternate(text):
    """
    return a string where every other word of text is repeated.

    Parameters
    ----------
    text : str
        Input string

    Returns
    -------
    Str
    """
    # ''.join([i + i if x%2 == 0 else i for x,i in enumerate
    # (re.findall(r'\S+\s?', text))]) works as well
    return re.sub(r'\s*(\b\S*\b)\s(\b\S*\b\s)', r'\1 \1 \2', text)
```

```python
text = ("Peter Piper picked a peck of pickled peppers "
        "A peck of pickled peppers Peter Piper picked "
        "If Peter Piper picked a peck of pickled peppers "
        "Where's the peck of pickled peppers Peter Piper picked ")
def repeat_alternate(text):
    return re.sub(r'\s*(\b\S+\b)\s(\b\S+\b\s)', r'\1 \1 \2', text)
```

<!-- grade_id: cell-e59aff47d6e9536d, points: 1, grade: true, solution: false -->
```python
text = ("Peter Piper picked a peck of pickled peppers "
        "A peck of pickled peppers Peter Piper picked "
        "If Peter Piper picked a peck of pickled peppers "
        "Where's the peck of pickled peppers Peter Piper picked ")
assert_equal(
    repeat_alternate(text),
    ("Peter Peter Piper picked picked a peck peck of pickled pickled peppers "
     "A A peck of of pickled peppers peppers Peter Piper Piper picked If If "
     "Peter Piper Piper picked a a peck of of pickled peppers peppers "
     "Where's the the peck of of pickled peppers peppers Peter Piper Piper "
     "picked ")
)
```

## Problem 5
Create a function `whats_on_the_bus` that will return the unique items that are on the bus according to the `text`.

<!-- grade_id: cell-c02c23356e069e11, points: 0, grade: false, solution: true -->
```python
def whats_on_the_bus(text):
    """
    Return the unique items that are on the bus according to the text.

    Parameters
    ----------
    text : str
        Input string

    Returns
    -------
    Set
    """
    return set(re.findall('The (\w+) on the bus', text))
```

Output:
```
14:33: W605 invalid escape sequence '\w'
```

```python
text = """
The wheels on the bus go round and round
Round and round, round and round
The wheels on the bus go round and round
All day long
The wipers on the bus go swish, swish, swish
Swish, swish, swish, swish, swish, swish
The wipers on the bus go swish, swish, swish
All day long
The horn on the bus goes beep, beep, beep
Beep, beep, beep, beep, beep, beep
The horn on the bus goes beep, beep, beep
All day long
The babies on the bus go wah, wah, wah
Wah, wah, wah, wah, wah, wah
The babies on the bus go wah, wah, wah
All day long
The wheels on the bus go round and round
Round and round, round and round
The wheels on the bus go round and round
All day long
"""
re.findall(r'The ([A-Za-z]+?) on the bus',text)
```

Output:
```
['wheels',
 'wheels',
 'wipers',
 'wipers',
 'horn',
 'horn',
 'babies',
 'babies',
 'wheels',
 'wheels']
```

<!-- grade_id: cell-2f60e46012eb1b21, points: 1, grade: true, solution: false -->
```python
text = """
The wheels on the bus go round and round
Round and round, round and round
The wheels on the bus go round and round
All day long
The wipers on the bus go swish, swish, swish
Swish, swish, swish, swish, swish, swish
The wipers on the bus go swish, swish, swish
All day long
The horn on the bus goes beep, beep, beep
Beep, beep, beep, beep, beep, beep
The horn on the bus goes beep, beep, beep
All day long
The babies on the bus go wah, wah, wah
Wah, wah, wah, wah, wah, wah
The babies on the bus go wah, wah, wah
All day long
The wheels on the bus go round and round
Round and round, round and round
The wheels on the bus go round and round
All day long
"""
items = whats_on_the_bus(text)
assert_equal(len(items), 4)
assert_equal(set(items), set(['babies', 'horn', 'wheels', 'wipers']))
```

## Problem 6
Create a function `to_list` that returns the list of items in `text` which were delimited by `,`, `+` or `and`.

<!-- grade_id: cell-5c113a9f3ffb394c, points: 0, grade: false, solution: true -->
```python
def to_list(text):
    """
    Return the unique items that are on the bus according to the text.

    Parameters
    ----------
    text : str
        Input string

    Returns
    -------
    Set
    """
    return re.split('[+,]|and', text)
```

<!-- grade_id: cell-607d1cbe582c1369, points: 1, grade: true, solution: false -->
```python
text = "a,b,candfoo bar+bazandd e+fee fi fo"
assert_equal(
    to_list(text), 
    ['a', 'b', 'c', 'foo bar', 'baz', 'd e', 'fee fi fo'])
```

Output:
```
3:19: W291 trailing whitespace
```

## Problem 7
Create a function `march_product` that returns the product of each `m by n` pair in `text`.

<!-- grade_id: cell-6fed3554cb468af7, points: 0, grade: false, solution: true -->
```python
def march_product(text):
    """
    Returns the product of each m by n pair in text

    Parameters
    ----------
    text : str
        Input string

    Returns
    -------
    List
    """
    pair = re.findall(r'(\d+) by (\d+)', text)
    return [int(p[0])*int(p[1]) for p in pair]
```

<!-- grade_id: cell-952d45adcb46ba4b, points: 1, grade: true, solution: false -->
```python
text = """
The ants go marching 1 by 1, hurrah, hurrah
The ants go marching 2 by 13, hurrah, hurrah
The ants go marching 42 by 8,
The little one stops to suck his thumb
And they all go marching down to the ground
To get out of the rain, BOOM! BOOM! BOOM!

The ants go marching 9 by 16, hurrah, hurrah
The ants go marching 54 by 7, hurrah, hurrah
The ants go marching 8 by 42,
The little one stops to tie his shoe
And they all go marching down to the ground
To get out of the rain, BOOM! BOOM! BOOM!
"""
assert_equal(march_product(text), [1, 26, 336, 144, 378, 336])
```

## Problem 8
Create a function `get_big` that will take in `items` and return the list of `ITEM`s that begin with `Big` but with an SKU that is not all numbers.

<!-- grade_id: cell-88383b660155ed69, points: 0, grade: false, solution: true -->
```python
def get_big(items):
    """
    Return the list of ITEMs that begin with Big but with an SKU that
    is not all numbers.

    Parameters
    ----------
    items : str
        Input string

    Returns
    -------
    List
    """
    matches = re.findall(r'\n\d*[A-Za-z]+\d*[A-Za-z]*\s(Big\s.*)', items)
    return matches
```

```python
items = """SKU ITEM
1A Big Red Box
A0 Big Bad Wolf
A1 Bigrams and Trigrams
02 Big Big World
BC Ain't Big Shoes
3C Bigger not Big"""
re.findall(r'^(?:[A-Z0-9]*(?:[A-Z]\d|\d[A-Z]|[A-Z])\S*?) (Big\b\s.*)', items, re.M)
```

Output:
```
['Big Red Box', 'Big Bad Wolf']
```

<!-- grade_id: cell-5efce4622970508a, points: 1, grade: true, solution: false -->
```python
items = """SKU ITEM
1A Big Red Box
A0 Big Bad Wolf
A1 Bigrams and Trigrams
02 Big Big World
BC Ain't Big Shoes
3C Bigger not Big"""

assert_equal(get_big(items), ['Big Red Box', 'Big Bad Wolf'])
```

## Problem 9
Create a function `find_chris` that accepts `names` as input and returns a list of all first names in `names` that contain the case-insensitive string `Chris` but with a last name that doesn't start with `B` or `M`.

<!-- grade_id: cell-c05a229e79e90f6f, points: 0, grade: false, solution: true -->
```python
def find_chris(test_str):
    """
    Accept names as input and return a list of all first names in names
    that contain the case-insensitive string Chris but with a last name
    that doesn't start with B or M

    Parameters
    ----------
    test_str : str
        Input string

    Returns
    -------
    List
    """
    return re.findall(r'(\w*chris\w*) [^BM]', names, re.I)
```

<!-- grade_id: cell-18375dd7af7df523, points: 1, grade: true, solution: false -->
```python
names = '''ID FIRST_NAME LAST_NAME
1 Christian Alis
A Chris Monterola
A1 Chris Hemsworth
02 Christian Bale
BC Christopher Nolan
3C Christopher de Leon
4F Marichris Chu'''

res = find_chris(names)
assert_equal(type(res), list)
assert_equal(set(res), 
             set(['Christian', 'Chris', 'Christopher', 'Christopher', 
                  'Marichris']))
```

Output:
```
12:23: W291 trailing whitespace
13:70: W291 trailing whitespace
```

## Problem 10
Create a function `get_client` that accepts `server_log` and returns a list of client IP , date/time of server access, and status code tuples from `log`. The value of `server_log` is shown below with relevant information highlighted in red.
<pre>
<font color="red">66.249.65.159</font> - - [<font color="red">06/May/2019:19:10:38 +0800</font>] "GET /news/53f8d72920ba2744fe873ebc.html HTTP/1.1" <font color="red">404</font> 177 "-" "Mozilla/5.0 (iPhone; CPU iPhone OS 6_0 like Mac OS X) AppleWebKit/536.26 (KHTML, like Gecko) Version/6.0 Mobile/10A5376e Safari/8536.25 (compatible; Googlebot/2.1; +http://www.google.com/bot.html)"
<font color="red">66.249.65.3</font> - - [<font color="red">06/May/2019:19:11:24 +0800</font>] "GET /?q=%E0%A6%AB%E0%A6%BE%E0%A7%9F%E0%A6%BE%E0%A6%B0 HTTP/1.1" <font color="red">200</font> 4223 "-" "Mozilla/5.0 (compatible; Googlebot/2.1; +http://www.google.com/bot.html)"
<font color="red">127.0.0.1</font> - - [<font color="red">06/May/2019:19:12:14 +0800</font>] "GET /?q=%E0%A6%A6%E0%A7%8B%E0%A7%9F%E0%A6%BE HTTP/1.1" <font color="red">200</font> 4356 "-" "Mozilla/5.0 (compatible; Googlebot/2.1; +http://www.google.com/bot.html)"
</pre>

<!-- grade_id: cell-75cc18ad07495089, points: 0, grade: false, solution: true -->
```python
def get_client(server_log):
    """
    Accept server_log and return a list of client IP , date/time of server
    access, and status code tuples from log.

    Parameters
    ----------
    server_log : str
        Input string

    Returns
    -------
    List
    """
    return re.findall(r'(\d+.\d+.\d+.\d+)\s-\s-\s\[(.*)\]\s\".*?\"\s(\d{3})',
                      server_log)
```

<!-- grade_id: cell-8fec960a9068214d, points: 1, grade: true, solution: false -->
```python
server_log = '''66.249.65.159 - - [06/May/2019:19:10:38 +0800] "GET /news/53f8d72920ba2744fe873ebc.html HTTP/1.1" 404 177 "-" "Mozilla/5.0 (iPhone; CPU iPhone OS 6_0 like Mac OS X) AppleWebKit/536.26 (KHTML, like Gecko) Version/6.0 Mobile/10A5376e Safari/8536.25 (compatible; Googlebot/2.1; +http://www.google.com/bot.html)"
66.249.65.3 - - [06/May/2019:19:11:24 +0800] "GET /?q=%E0%A6%AB%E0%A6%BE%E0%A7%9F%E0%A6%BE%E0%A6%B0 HTTP/1.1" 200 4223 "-" "Mozilla/5.0 (compatible; Googlebot/2.1; +http://www.google.com/bot.html)"
127.0.0.1 - - [06/May/2019:19:12:14 +0800] "GET /?q=%E0%A6%A6%E0%A7%8B%E0%A7%9F%E0%A6%BE HTTP/1.1" 200 4356 "-" "Mozilla/5.0 (compatible; Googlebot/2.1; +http://www.google.com/bot.html)"'''

assert_equal(get_client(server_log),
             [('66.249.65.159', '06/May/2019:19:10:38 +0800', '404'),
              ('66.249.65.3', '06/May/2019:19:11:24 +0800', '200'),
              ('127.0.0.1', '06/May/2019:19:12:14 +0800', '200')])
```

Output:
```
1:80: E501 line too long (324 > 79 characters)
2:80: E501 line too long (197 > 79 characters)
3:80: E501 line too long (189 > 79 characters)
```