```python
from numpy.testing import (assert_equal, assert_almost_equal, assert_raises, 
                           assert_raises_regex)
```

## Problem 1
Write a function `value_counts` that accepts a list `a_list` and an output filepath `out_path`. It saves the number of times each element in `a_list` appears to `out_path` as a comma-separated separated values file. Sort them from most to least frequent then by element if they have the same count. Don't include a header.

<!-- grade_id: cell-612f8d5c295e6aaa, points: 0, grade: false, solution: true -->
```python
def value_counts(a_list, out_path):
    """
    value_counts accepts a list a_list and an output filepath out_path. 
    It sorts the number of times each element in a_list appears from most to 
    least frequent then by element if they have the same count and saves it 
    to out_path as a comma-separated separated values file.
    
    Parameters
    ----------
    a_list : list
        List of elements that will be counted
    out_path : str
        csv file where the elements and their corresponding counts 
        will be saved
    """
    set_list = list(set(a_list))
    li_count = []   
    for sl in set_list:
        count = 0
        for al in a_list:
            if sl == al:
                count = count + 1
        li_count.append(count)

    li_zipped_items = list(zip(set_list,li_count))

#     output = sorted(li_zipped_items, key = lambda x:(x[0],x[1]))
    output = sorted(li_zipped_items, key = lambda x:(-x[1],x[0]))
    string = ''
    for out in output:

        string = string + str(out[0]) + "," + str(out[1]) + "\n"

    fp = out_path

    with open(fp, 'w') as f:
        f.write(string)
```

<!-- grade_id: cell-6741514de37183fb, points: 1, grade: true, solution: false -->
```python
!rm -f out_file.csv
value_counts(['a', 'b', 'c'], 'out_file.csv')
with open('out_file.csv') as f:
    assert_equal(
        f.read(),
        'a,1\n'
        'b,1\n'
        'c,1\n'
    )
!rm -f out_file.csv
value_counts(['a', 'b', 'c', 'b', 'a'], 'out_file.csv')
with open('out_file.csv') as f:
    assert_equal(
        f.read(),
        'a,2\n'
        'b,2\n'
        'c,1\n'
    )
```

```python
li_prac_hid = [('a',2),('b',1),('c',2)]
li_prac_1 = [('a',1),('b',1),('c',1)]
li_prac_2 = [('a',2),('b',2),('c',1)]

sorted(li_prac_1, key = lambda x:(-x[1]))

```

Output:
```
[('a', 1), ('b', 1), ('c', 1)]
```

## Problem 2
Create a function `is_subset` that checks whether `sublist` is a subset of `superlist`.

<!-- grade_id: cell-915b3dbfa6d30678, points: 0, grade: false, solution: true -->
```python
def is_subset(sublist, superlist, strict=True):
    """
    Check whether `sublist` is a subset of `superlist`
    
    Parameters
    ----------
    sublist : list
        List to check whether it is a subset of `superlist`
    superlist : list
        List to check whether `sublist` is one of its subsets
    strict : bool
        If `True`, the exact sequence of `sublist` must be found in 
        `superlist` for the `sublist` to be considered as a subset. If 
        `False`, `sublist` will be considered a subset of `superlist` as long
        as all members of `sublist` are found in `superlist`
        
    Returns
    -------
    is_subset : bool
        `True` is `sublist` is a subset of `superlist`, `False` otherwise
    """
    if strict:
    
        str_sub=''
        str_super = ''
        for sub_elem in sublist:
            str_sub = str_sub + str(sub_elem)
        str_sub

        for super_elem in superlist:
            str_super = str_super + str(super_elem)
        str_super
        
        output_true = str_sub in str_super
        return output_true
    else:
        test = [x for x in sublist if x in superlist]
        if test == sublist:
            return True
        else:
            return False
    
```

<!-- grade_id: cell-637e733c312fd315, points: 1, grade: true, solution: false -->
```python
assert_equal(is_subset([1, 2, 3], [0, 1, 2, 3, 4, 5]), True)
assert_equal(is_subset([1, 2, 3], [0, 1, 2, 4, 3, 5]), False)
assert_equal(is_subset([1, 2, 3], [0, 1, 2, 4, 3, 5], False), True)
assert_equal(is_subset([1, 'a', 3], [0, 1, 'a', 3, 4, 5]), True)
assert_equal(is_subset([1, 'a', 3], [0, 1, 'a', 4, 3, 5]), False)
assert_equal(is_subset([1, 'a', 3], [0, 1, 'a', 4, 3, 5], False), True)
```

```python

```

## Problem 3
Write a function called `has_duplicates` that takes a list and returns `True` if there is any element that appears more than once, `False` otherwise.

<!-- grade_id: cell-ada56c95461c367f, points: 0, grade: false, solution: true -->
```python
def has_duplicates(a_list):
    """
    has_duplicates takes a list and returns True if there is any element
    that appears more than once, False otherwise.
    
    Parameters
    ----------
    a_list : list
        List of elements that will be checked whether there are duplicates of 
        themselves within the list
        
    Returns
    -------
    output : bool
        True if there is any element that appears more than once, 
        False otherwise.
    """
    count = 0
    set_list = list(set(a_list))

    for set_i in set_list:
        for list_i in a_list:
            if list_i == set_i:
                count = count + 1
    if count > len(set_list):
        output = True
    else:
        output = False
    return output
```

<!-- grade_id: cell-5949ac9777599dbc, points: 1, grade: true, solution: false -->
```python
assert_equal(has_duplicates([1, 2, 3]), False)
assert_equal(has_duplicates([1, 2, 2, 3]), True)
assert_equal(has_duplicates([1, 2, 3, 2]), True)
```

## Problem 4
Write a function `count_words` that accepts the filepaths `input_file` and `output_file`, creates a dictionary where the key is a word and the value is the number of occurrences of that word in `input_file` then saves the dictionary into a pickle file `output_file`. A word is defined as a case-insensitive sequence of non-whitespace characters.

<!-- grade_id: cell-101ba68cff4103c9, points: 0, grade: false, solution: true -->
```python
def count_words(input_file, output_file):
    """
    count_words accepts the filepaths input_file and output_file, 
    creates a dictionary where the key is a word and the value is the number 
    of occurrences of that word in input_file then saves the dictionary into 
    a pickle file output_file. A word is defined as a case-insensitive 
    sequence of non-whitespace characters.
    
    Parameters
    ----------
    input_file : str
        filepath of file that will be read.
        
    output_file : str
        pickle file where the dictionary containing each word and its 
        corresponding value counts are its keys and values, respectively, 
        will be saved

    """
    with open(input_file) as f:
        text = f.read()
        
    new_text= text.replace('\n', ' ')
    spl = new_text.lower().split()
    set_spl = list(set(spl))

    wc ={}
    
    for set_i in set_spl:
        
            count = 0
            
            for spl_i in spl:

                if set_i == spl_i:
                    count = count + 1

            wc[set_i] = count
            
    with open(output_file, 'wb') as p:
        pickle.dump(wc,p)
```

<!-- grade_id: cell-6afa75204659d0ba, points: 1, grade: true, solution: false -->
```python
!rm -f word_count.pkl
import pickle
count_words('/mnt/data/public/gutenberg/README', 'word_count.pkl')
with open('word_count.pkl', 'rb') as f:
    wc = pickle.load(f)
assert_equal(len(wc), 78)
assert_equal(wc['this'], 5)
assert_equal(wc['1/2/7/1/12714'], 1)
```

## Problem 5 [2 pts]
Create a class `Person` that has the following methods:
* `move` that moves the person's position by the optional parameters `dx` (default: 0) and `dy` (default: 0)
* `get_position` that returns the person's current position as a tuple of x and y coordinates
* `is_infected` that returns `True` if the person is infected and `False` otherwise
* `set_infected` that assigns the person as infected
* `get_infected` that accepts another `Person` and becomes infected by that person if that person is infected and their Euclidean distance is less than the second argument `threshold`

`Person` can also be optionally initialized by a tuple of x and y coordinates, indicating the person's initial position (default: (0,0)). A `Person` is initially uninfected.

<!-- grade_id: cell-7363386ed5790b38, points: 0, grade: false, solution: true -->
```python
class Person():
    """
    The Person class contains 5 methods. The move method moves the position of
    the person. The get_position method returns the current coordinate 
    position of the person. The is_infected method returns True if the person
    is infected, and False otherwise. The set_infected method assigns the 
    person as infected. The get_infected method accepts another Person and
    becomes infected by that person if that person is infected and their
    Euclidean distance is less than the threshold.
    
    A person can be optionally initialized by a specific x and y coordinates, 
    otherwise the default position is set at (0,0). Lastly, the person is 
    initially uninfected
    
    """
    
    def __init__(self, init_pos=(0,0)):
        self.pos = init_pos
        self.infected = False
        
    def move(self, dx=0, dy=0):
        self.pos = (self.pos[0]+dx, self.pos[1]+dy)
        
    def get_position(self):
        return self.pos
    
    def set_infected(self):
        self.infected = True
    
    def is_infected(self):
        return self.infected
    
    def get_infected(self, stranger, threshold):
        
        euc_dist = (
            (stranger.pos[0]-self.pos[0])**2 \
            + (stranger.pos[1]-self.pos[1])**2
            )**0.5
        if stranger.infected == True and euc_dist < threshold:
            self.infected = True
```

<!-- grade_id: cell-07611009a85e2f77, points: 1, grade: true, solution: false -->
```python
person = Person()
assert_equal(person.get_position(), (0, 0))
assert_equal(person.is_infected(), False)

person.move(-1, 2)
assert_equal(person.get_position(), (-1, 2))
assert_equal(person.is_infected(), False)

person2 = Person((1, 2))
assert_equal(person2.get_position(), (1, 2))
assert_equal(person.is_infected(), False)

person2.move()
assert_equal(person2.get_position(), (1, 2))
assert_equal(person2.is_infected(), False)
```

<!-- grade_id: cell-7ee343c5ce1fa2b0, points: 1, grade: true, solution: false -->
```python
person1 = Person()
person2 = Person((1, 2))
person3 = Person((10, 20))
person3.set_infected()
person4 = Person((-1, -2))
person4.set_infected()

assert_equal(person1.is_infected(), False)
person1.get_infected(person2, 5)
assert_equal(person1.is_infected(), False)
person1.get_infected(person3, 5)
assert_equal(person1.is_infected(), False)
person1.get_infected(person4, 5)
assert_equal(person1.is_infected(), True)
```

## Problem 6
Create a class `QuarantinedPerson` that is a `Person` but whose position never changes regardless of the parameters passed to `move`.

<!-- grade_id: cell-41ba9afb8e1275e8, points: 0, grade: false, solution: true -->
```python
class QuarantinedPerson(Person):
    """
    The class QuarantinedPerson is a Person but whose position never changes 
    regardless of the parameters passed to move
    """
    
    def move(self, dx=0, dy=0):
        self.pos = self.pos
```

<!-- grade_id: cell-9a14040ff1676251, points: 1, grade: true, solution: false -->
```python
person1 = Person()
person2 = QuarantinedPerson((10, 0))
person2.set_infected()

assert_equal(person1.is_infected(), False)
person1.get_infected(person2, 5)
assert_equal(person1.is_infected(), False)
person2.move(-10)
person1.get_infected(person2, 5)
assert_equal(person1.is_infected(), False)
```

## Problem 7
Create a function `file_lines` that accepts an unlimited number of arguments having filepaths as values and returns a dictionary with the argument name as the key and the number of lines in the filepath as the value. It shouldn't include the argument if the filepath doesn't exist.

<!-- grade_id: cell-ef56a0ba3486faa3, points: 0, grade: false, solution: true -->
```python
def file_lines(**kwargs):
    """
    file_lines accepts an unlimited number of arguments having filepaths 
    as values and returns a dictionary with the argument name as the key and 
    the number of lines in the filepath as the value. It shouldn't include the
    argument if the filepath doesn't exist.
    
    Parameters
    ----------
    **kwargs : var = str
        var: where the filepath is assigned
        str: filepath of file that is to be read
        
    Returns
    -------
    d : dict
        Dictionary where the variable name is the key while the number of 
        lines is the values.
    """
    d = {}
    
    for key in kwargs:
        
        fp = kwargs[key]
        try:
            with open(fp) as f:
                text = f.read()
                count_lines = text.count('\n')
                
    
            d[key] = count_lines
        except:
            continue
    
    return d


```

<!-- grade_id: cell-34a26cef488d47b6, points: 1, grade: true, solution: false -->
```python
assert_equal(
    file_lines(g1996='/mnt/data/public/gutenberg/GUTINDEX.1996',
               g2020='/mnt/data/public/gutenberg/GUTINDEX.2020',
               gall='/mnt/data/public/gutenberg/GUTINDEX.ALL'),
    {'g1996': 1864, 'gall': 167706})
```

## Problem 8
Create a function `ten_div` that accepts a numerator and a denominator, and returns their quotient. The numerator can only be from 0 to 10, otherwise, it should raise a `TenDivError` which is a `ValueError`. It should return a `TenDivError` with message equal to `Error encountered: ` follow by the exception string if it encounters other exceptions.

<!-- grade_id: cell-804bdb56a0e766d8, points: 0, grade: false, solution: true -->
```python
class TenDivError(Exception):
    pass
def ten_div(num, denom):
    if not 0 <= num <= 10:
        message = "Error encountered: ValueError"
        raise TenDivError(message)
    elif denom == 0:
        message = 'Error encountered: division by zero'
        raise TenDivError(message)
    else:
        return num / denom
```

<!-- grade_id: cell-503e96dc6ad1a424, points: 1, grade: true, solution: false -->
```python
assert_almost_equal(ten_div(10, 2), 5.0)
assert_raises(TenDivError, ten_div, -2, 10)
assert_raises(TenDivError, ten_div, 11, 5)
assert_raises_regex(TenDivError, 'Error encountered: division by zero',
                    ten_div, 5, 0)
```

## Problem 9
Create a function `most_frequent` in a module `wordfreq` that accepts a filepath and returns the list of the nine most frequent words in the file sorted by decreasing frequency then by increasing alphabetical order. A word is case insensitive and defined as a sequence of non-whitespace characters.

<!-- grade_id: cell-ce4a1efe5626138d, points: 1, grade: true, solution: false -->
```python
import wordfreq
assert_equal(
    wordfreq.most_frequent('/mnt/data/public/gutenberg/README'),
    ['for', 'this', 'to', 'a', 'ebook', 'you', 'in', 'is', 'the']
)
```