# Applied Computational Statistics - Homework 0

This homework is a test.

Check out this cool equation:

\begin{align}
    F(x) &= \int_{-\infty}^{x} f(x) dx
\end{align}

```python
import numpy as np
```

In the cell below, create a function that adds two numbers.

<!-- grade_id: cell-666ccd8372a748dc, points: 0, grade: false, solution: true -->
```python
def add_numbers(a, b):
    '''
    this function adds two numbers
    '''
    return a+b
```

<!-- grade_id: cell-e2cf42bc57997f40, points: 0, grade: true, solution: false -->
```python
''' case 1 - add 1 + 1 '''
assert np.isclose(add_numbers(1, 1), 2)

```