```python
from numpy.testing import assert_equal
```

## Problem 1
Write a Python function that will count from 1 to an integer `n` (inclusive). For each count, print the counter value and a message separated by a single space. If the counter is even, the message should be `fizz`, `fuzz` if divisible by 3 and `foo` (only) if divisible by 6.

<!-- grade_id: cell-0e7266cd03645e41, points: 0, grade: false, solution: true -->
```python
def count_and_print(n):
    """Count from 1 to n whilst displaying a message
    
    Count from 1 to an integer n (inclusive). For each count, print the 
    counter value and a message separated by a single space. If the counter is
    even, the message should be `fizz`, `fuzz` if divisible by 3 and `foo` 
    (only) if divisible by 6.
    
    Parameters
    ----------
    n : integer
        Count up to this number (inclusive)
        
    Examples
    --------
    >>> count_and_print(12)
    1
    2 fizz
    3 fuzz
    4 fizz
    5
    6 foo
    7 
    8 fizz
    9 fuzz
    10 fizz
    11
    12 foo

    """
    
    for iter_p1 in range(1, n + 1):
        if iter_p1 % 6 == 0:
            counter = "foo"
            print(str(iter_p1) + " " + counter)
        elif iter_p1 % 3 == 0:
            counter = "fuzz"
            print(str(iter_p1) + " " + counter)
        elif iter_p1 % 2 == 0:
            counter = "fizz" 
            print(str(iter_p1) + " " + counter)
        else:
            print(str(iter_p1))
```

```python
%%capture out1
count_and_print(12)
```

```python
%%capture out2
count_and_print(18)
```

<!-- grade_id: cell-5e68bd9f3673ce01, points: 1, grade: true, solution: false -->
```python
assert_equal(
    out1.stdout,
    '1\n'
    '2 fizz\n'
    '3 fuzz\n'
    '4 fizz\n'
    '5\n'
    '6 foo\n'
    '7\n'
    '8 fizz\n'
    '9 fuzz\n'
    '10 fizz\n'
    '11\n'
    '12 foo\n')
assert_equal(
    out2.stdout,
    '1\n'
    '2 fizz\n'
    '3 fuzz\n'
    '4 fizz\n'
    '5\n'
    '6 foo\n'
    '7\n'
    '8 fizz\n'
    '9 fuzz\n'
    '10 fizz\n'
    '11\n'
    '12 foo\n'
    '13\n'
    '14 fizz\n'
    '15 fuzz\n'
    '16 fizz\n'
    '17\n'
    '18 foo\n'
)
```

## Problem 2
Write a function `overlap_interval` that accepts parameters `start1`, `end1`, `start2` and `end2` corresponding to the endpoints of two numeric intervals. It should print the endpoints of their overlaps or print the message `No overlap!` if they don't overlap.

<!-- grade_id: cell-c3a64bdf8b11fb95, points: 0, grade: false, solution: true -->
```python

def overlap_interval(start1,end1,start2,end2):
    """
    overlap_interval() accepts parameters start1, end1, start2 and end2 
    corresponding to the endpoints of two numeric intervals. 
    It should print the endpoints of their overlaps 
    or print the message No overlap! if they don't overlap.
    
    Parameters
    ----------
    start1 : integer
        starting point for first numeric interval
    end1 : integer
        end point for first numeric interval that is less than start1
    start2 : integer
        starting point for second numeric interval
    end2 : integer
        end point for second numeric interval that is less than start2
    
    Returns
    -------
    "invalid intervals" : str
        if parameters are invalid
    PO : str
        point of overlap if intervals overlap at exactly one point
    EPOL : str
        end points of the overlap of the intervals
    
    """
    interval_1 = range(start1,end1+1)
    interval_2 = range(start2,end2+1)
    
    list_of_overlap_vals = [i for i in interval_1 if i in interval_2]
    
    if start1 > end1 or start2 > end2:
        return "invalid intervals"
    
    elif len(list_of_overlap_vals) == 1:
        PO = print(str(list_of_overlap_vals[0]))
        return PO
    elif len(list_of_overlap_vals) != 0:
        EPOL = print(str(list_of_overlap_vals[0]) + " " + \
                     str(list_of_overlap_vals[-1]))
        return EPOL
    else:
        return print("No overlap!")
        
```

```python
%%capture out1
overlap_interval(-5, 10, 5, 12)
```

```python
%%capture out2
overlap_interval(-5, 10, 11, 12)
```

```python
%%capture out3
overlap_interval(-5, 10, -10, 12)
```

```python
%%capture out4
overlap_interval(5, 8, -10, 12)
```

```python
%%capture out5
overlap_interval(-12, 12, 8, 10)
```

<!-- grade_id: cell-7812ed25101d9ddf, points: 1, grade: true, solution: false -->
```python
assert_equal(out1.stdout, '5 10\n')
assert_equal(out2.stdout, 'No overlap!\n')
assert_equal(out3.stdout, '-5 10\n')
```

## Problem 3
Write a function that returns the cosine of an input (in radians) using [Maclaurin's approximation](https://en.wikibooks.org/wiki/Trigonometry/Power_Series_for_Cosine_and_Sine). It should use a `while` loop to compute terms of the summation until the last term is smaller than $10^{-15}$.

<!-- grade_id: cell-54183911ca9ecc32, points: 1, grade: true, solution: true -->
```python

def cosine(theta):
    """Return cosine of theta using Maclaurin's approximation
    
    Parameters
    ----------
    theta : float
        Value to compute cosine of, in radians
        
    Returns
    -------
    cos_theta : float
        Cosine of theta
    """
    

    k=0
    
    def factorial_(fact_arg):
        total_fac = 1
        for iter_fac in range(1,fact_arg+1):
            total_fac = total_fac * iter_fac
        return total_fac
    answer_p3 = 0 
    
    term = 1
    
    
    while abs(term) > 1e-15:
        
        term = ((-1)**(k) / (factorial_(2 * k)))*(theta**((2 * k)))
        
        answer_p3 = answer_p3 + term
        k = k + 1
    print(f"final answer : {answer_p3:.1f}")
    print("k:",k)
        
cosine(2*(3.14))
```

Output:
```
final answer : 1.0
k: 21
```

## Problem 4
The greatest common divisor (GCD) of $a$ and $b$ is the largest number that divides both of them with no remainder.

One way to find the GCD of two numbers is based on the observation that if $r$ is the remainder when $a$ is divided by $b$, then $gcd(a, b) = gcd(b, r)$. As a base case, we can use $gcd(a, 0) = a$.

Write a function called `gcd` that takes parameters `a` and `b` and returns their greatest common divisor.

<!-- grade_id: cell-c29ca12ec75789d2, points: 0, grade: false, solution: true -->
```python
def gcd(a,b):
    """
    gcd() takes parameters a and b and returns their greatest common divisor.
    
    Parameters
    ----------
    a : int
    b : int
    
    Returns
    -------
    gcd : int
        Greatest Common Divisor of a and b    
    """
    range_a = range(1,a+1)
    range_b = range(1,b+1)
    divisor_list_a = []
    divisor_list_b = []
    for a_iter in range_a:
        if a%a_iter == 0 :
            divisor_list_a.append(a_iter)
    for b_iter in range_b:
        if b%b_iter == 0 :
            divisor_list_b.append(b_iter)
    
    com_div = [i for i in divisor_list_a if i in divisor_list_b]

    gcd = max(com_div)
    return gcd
```

<!-- grade_id: cell-2f070dd683636894, points: 1, grade: true, solution: false -->
```python
assert_equal(gcd(4, 2), 2)
assert_equal(gcd(2, 4), 2)
assert_equal(gcd(13, 5), 1)
assert_equal(gcd(20, 24), 4)
```

## Problem 5
Write a function `biased_sum` that returns the sum of positive integers passed to it but numbers that are multiples of the `base` are doubled before being added. By default, `base` is 2. Return `None` if there are no numbers passed to it.

<!-- grade_id: cell-6295767c33976fea, points: 0, grade: false, solution: true -->
```python
def biased_sum(*nums, base=2):
    """
    biased_sum() returns the sum of positive integers passed to it but numbers
    that are multiples of the base are doubled before being added. 
    By default, base is 2. Return None if there are no numbers passed to it.
    
    Parameters
    ----------
    *nums : positive int
        values to be summed  
    base : int
        nums will be doubled before adding if it 
        is a multiple of this value.
    
    Returns
    -------
    total : positive int
        sum of nums parameter
    None : NoneType
        if no numbers are passed
    
    """
    total = 0
    if len(nums) != 0:
        for iter_p5 in nums:
            if iter_p5%base == 0:
                iter_p5 = iter_p5*2
                total = total +iter_p5
            else:
                total = total + iter_p5
        return total
    else:
        
        return print(None) 
```

<!-- grade_id: cell-269f051cdd06af87, points: 1, grade: true, solution: false -->
```python
assert_equal(biased_sum(1, 2, 3), 8)
assert_equal(biased_sum(1, 2, 3, 4, 5, 6, 3), 36)
assert_equal(biased_sum(1, 2, 3, 4, 5, 6, base=3), 30)
```

## Problem 6
Write a function `last_in_sequence` that accepts a string of digits, which can have a repeating sequence of 0 to 9 in it but may have intervening digits between them. The function should return the last digit in the sequence found in the string. It should return `None` if there is no such sequence in the string.

Examples:
* **0**9**12**251**3**2**4**21 $\rightarrow$ 4
* 1**0**9**12**251**3**2**4**21 $\rightarrow$ 4
* **012345678901234** $\rightarrow$ 4
* 1912251324 $\rightarrow$ `None`

<!-- grade_id: cell-f26b1b122b693409, points: 0, grade: false, solution: true -->
```python
def last_in_sequence(string_p6):
    """
    last_in_sequence() accepts a string of digits, which can have a repeating 
    sequence of 0 to 9 in it but may have intervening digits between them. 
    The function should return the last digit in the 
    sequence found in the string. It should return None if there is no such 
    sequence in the string.
    
    Parameters
    ----------
    string_p6 : str
        sequence of string digits  
        
    Returns
    -------
    int(string_p6[-1]) : int
        last digit of the string_p6 parameter
        
    None : NoneType
        if string_p6 does not contain 0 and 9
    
    """
    set_str = set(string_p6)
    sorted_set_str = sorted(set_str)
    
    if sorted_set_str[0] == '0' and sorted_set_str[-1] == '9':
        return int(string_p6[-1])
    else:
        return None

```

<!-- grade_id: cell-19149d85bde88dea, points: 1, grade: true, solution: false -->
```python
assert_equal(last_in_sequence('0912251324'), 4)
assert_equal(last_in_sequence('10912251324'), 4)
assert_equal(last_in_sequence('012345678901234'), 4)
assert_equal(last_in_sequence('1912251324'), None)
```

## Problem 7
Write a function `check_password` that returns `True` if the input password satisfies the following criteria:

* at least 8 characters in length
* contains at least one lowercase letter
* contains at least one uppercase letter
* contains at least one number

It should return `False` if at least one of the criteria is not satisfied.

<!-- grade_id: cell-a72d390a2aa6048b, points: 0, grade: false, solution: true -->
```python
def check_password(password):
    """
    check_password() returns True if the input password satisfies 
    the following criteria:

    - at least 8 characters in length
    - contains at least one lowercase letter
    - contains at least one uppercase letter
    - contains at least one number
    
    It should return False if at least one of the criteria is not satisfied.
    
    Parameters
    ----------
    password : str
        sequence of strings  
        
    Returns
    -------
    False : bool 
        if at least one criteria is violated
        
    True : bool
        if all criteria are satisfied
    
    """
    count_lower = 0
    count_upper = 0
    count_int = 0
    for lowerc_iter in password:
        if lowerc_iter.islower() == True:
            count_lower = count_lower + 1
        else:
            if lowerc_iter.isupper() == True:
                count_upper = count_upper+1
            else:
                count_int = count_int + 1
    if len(password) < 8:
        return False
    elif count_lower >= 1 and count_upper >= 1 and count_int >= 1:
        return True
    else:
        return False
```

<!-- grade_id: cell-b8041b3660eec0bd, points: 1, grade: true, solution: false -->
```python
assert_equal(check_password('abcABC1'), False)
assert_equal(check_password('abcABC12'), True)
assert_equal(check_password('abcABCDE'), False)
assert_equal(check_password('abc12345'), False)
assert_equal(check_password('12345ABC'), False)
```

## Problem 8
Write a function `is_palindrome` that returns `True` if the input string is a palindrome and `False` otherwise.

<!-- grade_id: cell-c7e575e3d909fd56, points: 0, grade: false, solution: true -->
```python
def is_palindrome(text):
    """" 
    is_palindrome() returns True if the input string is a palindrome
    and False otherwise. 
    
     Parameters
    ----------
    text : str
        sequence of strings  
        
    Returns
    -------
    False : bool 
        if text parameter is not a palindrome
        
    True : bool
         if text parameter not a palindrome
    """
    li_no_space = text.split(" ")
    word=""
    for iter_p8 in li_no_space:
        word = word + iter_p8
    processed_text = word
    
    
    if processed_text[::-1] == processed_text:
        return True

    else:
        return False
is_palindrome(' ab ba')
```

Output:
```
True
```

<!-- grade_id: cell-71c90ba9148f6663, points: 1, grade: true, solution: false -->
```python
assert_equal(is_palindrome('abcd'), False)
assert_equal(is_palindrome('abba'), True)
assert_equal(is_palindrome('abeba'), True)
assert_equal(is_palindrome('abcda'), False)
assert_equal(is_palindrome('ab cd'), False)
assert_equal(is_palindrome('ab ba'), True)
assert_equal(is_palindrome(' ab ba'), True)
assert_equal(is_palindrome('ab  ba'), True)
```

## Problem 9
Write a function `create_squares` that takes in the number of stars on a side of a square then returns a 2$\times$2 grid of squares as a string. Separate each line by the newline (`\n`) character.

<!-- grade_id: cell-9d0bd755f6457ac0, points: 0, grade: false, solution: true -->
```python
def create_squares(num_stars):
    """
    create_squares() takes in the number of stars on a side of a square then returns
    a 2 × 2 grid of squares as a string. Separate each line by
    the newline (\n) character.
    
     Parameters
    ----------
    num_stars : positive int
        number of stars inside the top left and bottom right squares of a grid 
        
    Returns
    -------
    output : str 
        grid with squares filled with num_stars amount of stars on the 
        top left and bottom right squares while empty on the remaining squares
    
    """
    h_border = "- "*num_stars
    filling = "* "*num_stars
    no_fill = "  "*num_stars

    row_1 = "+ " + h_border + "+ " + h_border + "+\n"
    left_fill_row = "| " + filling + "| " + no_fill + "|\n"
    right_fill_row = "| " + no_fill + "| " + filling + "|\n"


    output = (row_1 + left_fill_row*num_stars + row_1 + 
              right_fill_row*num_stars + row_1) 

    return output
```

<!-- grade_id: cell-1f0947cbeb45f780, points: 1, grade: true, solution: false -->
```python
assert_equal(
    create_squares(1),
"""+ - + - +
| * |   |
+ - + - +
|   | * |
+ - + - +
"""
)
assert_equal(
    create_squares(3),
"""+ - - - + - - - +
| * * * |       |
| * * * |       |
| * * * |       |
+ - - - + - - - +
|       | * * * |
|       | * * * |
|       | * * * |
+ - - - + - - - +
"""
)
```

## Problem 10
Write a function `create_grid` that takes in the number squares on a side and the number of stars on a side of a square then returns the grid of squares as a string. Separate each line by the newline (`\n`) character.


<!-- grade_id: cell-d554859dd48da4e8, points: 0, grade: false, solution: true -->
```python

def create_grid(num_squares, num_stars):
    """
    create_grid takes in the number squares on a side and the number of stars 
    on a side of a square then returns the grid of squares as a string. 
    Separate each line by the newline (\n) character.
    
    Parameters
    ----------
    num_stars : positive int
        number of stars inside alternating squares of the grid 
        
    num_squares : positive int
        number of squares on a side of the grid
        
    Returns
    -------
    grid_string : str 
        grid with alternating squares filled with num_stars amount of stars 
        on the top left and bottom right squares while 
        empty on the remaining squares
    """

    grid_string = ("+ " + "- " * num_stars) * num_squares + "+\n"

    for k in range(1, num_squares + 1):
        if k % 2 == 0:
            for i in range(1, num_stars + 1): 
                for j in range(1, num_squares + 1):
                    if j % 2 == 0:
                        grid_string += str("| " + "* " * num_stars)
                    else:
                        grid_string += str("| " + "  " * num_stars)

                grid_string += "|\n"
        else:
            for i in range(1, num_stars + 1): 
                for j in range(1, num_squares + 1):
                    if j % 2 == 0:
                        grid_string += str("| " + "  " * num_stars)
                    else:
                        grid_string += str("| " + "* " * num_stars)

                grid_string += "|\n"

        grid_string += str("+ " + "- " * num_stars) * num_squares + "+\n"

    return grid_string
```

<!-- grade_id: cell-804b324f525cec69, points: 1, grade: true, solution: false -->
```python
assert_equal(
    create_grid(1, 1),
"""+ - +
| * |
+ - +
""")
assert_equal(
    create_grid(2, 4),
"""+ - - - - + - - - - +
| * * * * |         |
| * * * * |         |
| * * * * |         |
| * * * * |         |
+ - - - - + - - - - +
|         | * * * * |
|         | * * * * |
|         | * * * * |
|         | * * * * |
+ - - - - + - - - - +
""")
assert_equal(
    create_grid(4, 2),
"""+ - - + - - + - - + - - +
| * * |     | * * |     |
| * * |     | * * |     |
+ - - + - - + - - + - - +
|     | * * |     | * * |
|     | * * |     | * * |
+ - - + - - + - - + - - +
| * * |     | * * |     |
| * * |     | * * |     |
+ - - + - - + - - + - - +
|     | * * |     | * * |
|     | * * |     | * * |
+ - - + - - + - - + - - +
""")
```