```python
from numpy.testing import assert_equal
```

## Problem 1
Write a function called `chop` that takes a list, modifies it by removing the first and last elements, and returns `None`.

<!-- grade_id: cell-84f06a445961e437, points: 0, grade: false, solution: true -->
```python
def chop(a_list):
    """ takes a list, modifies it by removing the first and last elements, 
    and returns None
    
    Parameters
    ----------
    a_list : list
        list of any element type
        
    Returns
    -------
    None : NoneType
    """
    for i in [0, -1]:
        a_list.pop(i)
    return None
```

<!-- grade_id: cell-645111f6cb08b363, points: 1, grade: true, solution: false -->
```python
a = [1, 2, 3, 4, 5]
assert chop(a) is None
assert_equal(a, [2, 3, 4])
```

## Problem 2
Write a function `sum_multiples` that computes the sum of multiples of 3 or 5 in an input list.

<!-- grade_id: cell-68e647afb7616bc5, points: 0, grade: false, solution: true -->
```python
def sum_multiples(a_list):
    """computes the sum of multiples of 3 or 5 in an input list
     
    Parameters
    ----------
    a_list : list
        list of integers
        
    Returns
    -------
    sum_p2 : int
        sum of multiples of 3 or 5 in elements of a_list
     """
        
    li_mul_3 = [i for i in a_list if i%3 == 0]
    li_mul_5 = [j for j in a_list if j%5 == 0]
    sum_p2 = sum(li_mul_3 + li_mul_5)
    return sum_p2
```

<!-- grade_id: cell-ba33eaa3bd8ee05a, points: 1, grade: true, solution: false -->
```python
assert_equal(sum_multiples(list(range(10))), 23)
assert_equal(sum_multiples(list(range(1, 10, 2))), 17)
assert_equal(sum_multiples(list(range(1, 10, 3))), 0)
```

## Problem 3
Write a function `rotate` that rotates a list by `k` elements.

<!-- grade_id: cell-7be8f562c1868e66, points: 0, grade: false, solution: true -->
```python
def rotate(numbers, k):
    """Rotate numbers by k elements
    Parameters
    ----------
    numbers : list
        list of integers
    k : int
        degree of rotation of elements of list 
        
    Returns
    -------
    rotated_list : list
        rotated list of original list
    """

    rotated_list = [0]*len(numbers)
    
    thresh = len(numbers) - k
    for i_p3 in range(len(numbers)):
        if i_p3 <= thresh - 1:
            rotated_list[i_p3+k] = numbers[i_p3]
        else:
            rotated_list[i_p3-thresh] = numbers[i_p3]
    return rotated_list    
    
```

<!-- grade_id: cell-8a8f9e9053ad1d3c, points: 1, grade: true, solution: false -->
```python
assert_equal(rotate([1,2,3,4,5,6], 4), [3,4,5,6,1,2])
assert_equal(rotate([1,2,3,4,5,6], -4), [5,6,1,2,3,4])
```

## Problem 4
Write a function `on_all` that applies a function to every element of a list.

<!-- grade_id: cell-8205a5864316caaf, points: 0, grade: false, solution: true -->
```python
def on_all(func, a_list):
    """applies a function to every element of a list.
    Parameters
    ----------
    func : function
        function to be applied to every element of a_list
    a_list : list
        list of integers 
        
    Returns
    -------
    output_p4 : list
        list of elements after func application
     """
    output_p4 = [func(i_p4) for i_p4 in a_list]
    
    return output_p4
```

<!-- grade_id: cell-9c4be7904b0af281, points: 1, grade: true, solution: false -->
```python
assert_equal(on_all(abs, [1, -2, 3, -4, 5]), [1, 2, 3, 4, 5])
assert_equal(on_all(lambda x: x**2, [1, -2, 3, -4, 5]), 
                   [1, 4, 9, 16, 25])
```

## Problem 5
Write a function `matrix_times_vector` that accepts matrix `mat` as a list of lists and vector `vec` as a list, and returns their product as a list. Assume the input dimensions are compatible.

<!-- grade_id: cell-51f7405f0170c1dd, points: 0, grade: false, solution: true -->
```python
def matrix_times_vector(mat, vec):
    """accepts matrix mat as a list of lists and vector vec as a list, 
    and returns their product as a list. Assume the input dimensions 
    are compatible.
    
    Parameters
    ----------
    mat : list
        list of lists that represent a matrix
    vec : list
        list that represents a vector 
        
    Returns
    -------
    prod : list
        product of mat and vec as a list
    """
    mat_size = [len(mat),len(mat[0])]
    vec_size = [len(vec),1]

    prod_size =[mat_size[0],vec_size[1]]

    prod = [0]*prod_size[0]

    for i_row in range(mat_size[0]):
        for i_col in range(mat_size[1]):
            prod[i_row] = prod[i_row] + mat1[i_row][i_col] * vec[i_col]
    return prod 
```

<!-- grade_id: cell-679aba3b9570c7c0, points: 1, grade: true, solution: false -->
```python
assert_equal(
    matrix_times_vector(
        [[1, 2, 3],
         [7, 6, 8]],
        [9, 10, 11]
    ),
    [62, 211]
)
```

## Problem 6
Write a function `coder` that automatically converts alphanumeric English text to Morse code and vice versa. Empty `text` should return an empty string. Assume the message consists only of alphanumeric characters and space.

<!-- grade_id: cell-7cd6dfd0b24acd83, points: 0, grade: false, solution: true -->
```python
def coder(text, to_morse=True):
    """Convert English text to Morse code and vice versa
    
    Parameters
    ----------
    text : str
        English text or Morse code text
    to_morse : boolean
        Converts text to Morse code if `True`, to English otherwise
        
    Returns
    -------
    coded_text : str
        Text converted to Morse code or English
    """
    morse = ['.-','-...','-.-.','-..','.','..-.','--.','....','..','.---',
             '-.-','.-..','--','-.','---','.--.','--.-','.-.','...','-','..-',
             '...-','.--','-..-','-.--','--..','.----','..---','...--',
             '....-','.....','-....','--...','---..','----.','-----', ' ']

    alphabet = ['a','b','c','d','e','f','g','h','i','j','k','l','m','n',
                'o','p','q','r','s','t','u','v','w','x','y','z','1','2',
                '3','4','5','6','7','8','9','0',' ']


    zip_alpha_morse = zip(alphabet,morse)

    dic = {i:j for i,j in zip_alpha_morse}

    
    trans_alpha_morse = ''
    trans_morse_alpha_sentence = []
    
    if to_morse is False:
        morse_words = text.lower().split('   ')
        
        for mw_i in morse_words:
            trans_morse_alpha_word = ''
            
            for char in mw_i.split(' '):
                #morse to alpha
                list_keys = list(dic.keys())
                list_values = list(dic.values())

                index_val = list_values.index(char)
                corr_key = list_keys[index_val]
                trans_morse_alpha_word = trans_morse_alpha_word + corr_key
                
            trans_morse_alpha_sentence.append(trans_morse_alpha_word)
            
        morse_alpha_output = ' '.join(trans_morse_alpha_sentence)        
        return morse_alpha_output.upper()
    
    elif to_morse:
        for char in ' '.join(text.lower()):
            #alpha to morse
            trans_alpha_morse = trans_alpha_morse + dic[char]
        return trans_alpha_morse
    else:
            return ''
```

<!-- grade_id: cell-dc1c5996cd5902b2, points: 1, grade: true, solution: false -->
```python
assert_equal(
    coder('Hello world'), 
    '.... . .-.. .-.. ---   .-- --- .-. .-.. -..'
)
assert_equal(
    coder('-.. .- - .-   ... -.-. .. . -. -.-. .', False),
    'DATA SCIENCE'
)
```

## Problem 7
Complete the function below.

<!-- grade_id: cell-ac1556c1573c42d3, points: 0, grade: false, solution: true -->
```python
def sort_by_key(items_with_keys, ascending=True):
    """Sort items_with_keys based on the keys then by item for same keys
    
    Parameters
    ----------
    items_with_keys : list of tuples
        list of (item, key) tuples to be sorted
    ascending : bool
        Sort in ascending order if True, in descending order otherwise
    """
    lis = [('B', 10), ('A', 12), ('C', 3), ('D', 5)]

    dict_p7 = {i[1]:i[0] for i in lis}
    
    if ascending:
        sorted_keys = sorted(dict_p7)
        
    else:
        sorted_keys = sorted(dict_p7 , reverse=True)
    sorted_items = [dict_p7[i] for i in sorted_keys]
    sorted_output = list(zip(sorted_items,sorted_keys))
    
    return sorted_output
```

<!-- grade_id: cell-da5dfd9915098c50, points: 1, grade: true, solution: false -->
```python
assert_equal(
    sort_by_key([('B', 10), ('A', 12), ('C', 3), ('D', 5)]),
    [('C', 3), ('D', 5), ('B', 10), ('A', 12)])
assert_equal(
    sort_by_key([('B', 10), ('A', 12), ('C', 3), ('D', 5)], False),
    [('A', 12), ('B', 10), ('D', 5), ('C', 3)])
```

## Problem 8
Write a function `count_words` that accepts a string `text` and returns a dictionary where the key is a word and the value is the number of occurrences of that word in `text`. A word is defined as a case-insensitive sequence of non-whitespace characters.

<!-- grade_id: cell-bc907ec0368f8aaf, points: 0, grade: false, solution: true -->
```python
def count_words(text):
    """accepts a string text and returns a dictionary where the key is a 
    word and the value is the number of occurrences of that word in text. 
    A word is defined as a case-insensitive sequence of 
    non-whitespace characters.
     
    Parameters
    ----------
    text : str
        sequence of words
        
    Returns
    -------
    wc : dictionary
        words for keys and word counts for values
    """
    #from collections import Counter
    for i in range(len(text)):
        new_text = text.replace('\n',' ')
        spl = new_text.lower().split()
        set_spl = list(set(spl))

    wc = {}


    for set_i in set_spl:
        count = 0
        for spl_i in spl:

            if set_i == spl_i:
                count = count + 1

        wc[set_i] = count
        

    #return Counter(spl)

    return wc
```

<!-- grade_id: cell-6e8fc7aff97f3758, points: 1, grade: true, solution: false -->
```python
wc = count_words("""
Mary had a little lamb
Little lamb little lamb
Mary had a little lamb
It's fleece was white as snow

Everywhere that Mary went 
Mary went Mary went
Everywhere that Mary went
The lamb was sure to go
""")
assert_equal(len(wc), 18)
assert_equal(wc['lamb'], 5)
assert_equal(wc["it's"], 1)
```

## Problem 9

Write a function `display_tree` that accepts a dictionary whose values can be numbers, strings or dictionaries, and returns a string. The contents of the returned string is as follows:
* For every key-value pair in the dictionary, sort by key then show the key, one per line, followed by a colon (:)
* If the value is not a dictionary, show the value prefixed by a space
* If the value is a dictionary, display the contents on the dictionary from the next line, following the rules in this list
* Indent by two spaces for every level of the dictionary
* Use the newline character (`\n`) for creating lines

<!-- grade_id: cell-233f3054b37827fc, points: 0, grade: false, solution: true -->
```python
def display_tree(dict_):
    """accepts a dictionary whose values can be numbers, 
    strings or dictionaries, and returns a string. The contents of the 
    returned string is as follows:

    -For every key-value pair in the dictionary, sort by key then show the key,
    one per line, followed by a colon (:)
    
    -If the value is not a dictionary, show the value prefixed by a space
    
    -If the value is a dictionary, display the contents on the dictionary from 
    the next line, following the rules in this list
    
    -Indent by two spaces for every level of the dictionary
    
    -Use the newline character (\n) for creating lines
    
    Parameters
    ----------
    dict_ : dict
        dictionary whose values can be numbers, strings, or dictionaries.
        
    Returns
    -------
    level(dict_) : str
        contents of the dictionary following the rules above.
    
    """
    def level(dict_, lev=0):
        keys = dict_.keys()
        key_types = [type(dict_[i]) for i in keys]

        sorted_keys = sorted(keys)
        output = ''

        for srt_ks in sorted_keys:

            if type(dict_[srt_ks]) is dict:
                
                    
                output = output+ '  '*lev  + str(srt_ks) +":\n" + level(dict_[srt_ks], lev=lev+1)
            else: 
                 
                output =  output+ '  '*lev  + str(srt_ks) + ': ' + str(dict_[srt_ks]) + '\n'
                
                
        return output
    
    return level(dict_)
```

<!-- grade_id: cell-63304af010a78f81, points: 1, grade: true, solution: false -->
```python
assert_equal(
    display_tree({
    'foo': 1,
    'bar': {
        'fee': 'eenie',
        'fi': 'meenie',
        'fo': 'minnie',
        'fum': {
            'mo': 'lover'
        },
    'baz':
        {
            1: 'potato',
            2: 'potatoes',
            3: 'potatoes'
        }
    }
    }),
    """bar:
  baz:
    1: potato
    2: potatoes
    3: potatoes
  fee: eenie
  fi: meenie
  fo: minnie
  fum:
    mo: lover
foo: 1
"""
)
```

## Problem 10

Write a function `get_nested_key_value` that accepts a nested dictionary `nested_dict` and a string `key`, which corresponds to the key at each level joined by a period (.). Missing keys should return `None`. Assume that keys are alphanumeric.

<!-- grade_id: cell-276faaed0193b2c3, points: 0, grade: false, solution: true -->
```python
def get_nested_key_value(nested_dict,key):
    """ accepts a nested dictionary nested_dict and a string key, 
    which corresponds to the key at each level joined by a period (.). 
    Missing keys should return None. Assume that keys are alphanumeric.
    
    Parameters
    ----------
    nested_dict : dict
        dictionary with nested dictionaries
        
    Returns
    -------
    key : str
        string that corresponds to the key at each level joined by 
        a period (.)
    """
    spl_key = key.split('.')
    
    
    cumul = nested_dict
    
    for i in spl_key:
        if i in cumul:
            cumul = cumul[i]
        else:
            cumul = None
            break
            
            
    return cumul
```

<!-- grade_id: cell-99592c2742ae0394, points: 1, grade: true, solution: false -->
```python
d = {
    'a': {
        'a1': {
            'aa1': 'bar',
            'aa2': 'baz'
        },
        'a2': 'bazbar',
        'a3': 0
    },
    'b': 1,
    'c': {
        'c1': 'eenie',
        'c2': {
            'x': 'meenie',
            'y': 'minnie',
            'z': 'mo'
        },
        'c3': 'lover'
    }
}
assert_equal(
    get_nested_key_value(d, 'a.a1.aa2'),
    'baz'
)
assert_equal(
    get_nested_key_value(d, 'a.a1'),
    {
        'aa1': 'bar',
        'aa2': 'baz'
    }
)
assert_equal(
    get_nested_key_value(d, 'a.a1.aa2.aaa1'),
    None
)
assert_equal(
    get_nested_key_value(d, 'a.a1.aa5'),
    None
)
```